---
title: "Lesson 10: Simulations" 
author: "Copyright 2025 Psychology Department. Hebrew University of Jerusalem. All rights reserved"

output:
  learnr::tutorial:
    progressive: true
    allow_skip: false

description: >
 <div style='direction: rtl;'>
 בשיעור זה נלמד לערוך סימולציות בR לטובת בניית התפלגויות דגימה וחישוב עוצמת מבחן.
 </div>
       
runtime: shiny_prerendered  
editor_options: 
  markdown: 
    wrap: 72
---

<style>
h1, h2, h3, h4, h5, h6 {
  direction: rtl;
}
p {
  direction: rtl;
}
.text-block1 {
  direction: rtl;       /* Set text direction to right-to-left */
  text-align: right;
  background-color: #e7f3fe; /* Light grey background */
  padding: 10px;
  border-radius: 5px;
  border: 1px solid #ddd; /* Light border */
  margin: 10px 0; /* Space around the block */
}
</style>


```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(Rcourse)
library(ggplot2)
library(tidyverse)

knitr::opts_chunk$set(error = TRUE)
knitr::opts_chunk$set(warning = FALSE)

gradethis::gradethis_setup()

set.seed(123456)


{
exercises_df = data.frame(exercises = c("q1", "q2a_hanna", 
                                        "ex1_normal_p_value",
                                        "q3_sampling_dist_normal",
                                        "analytic_sample_dist_exercise","analytic_s_dist_exercise2",
                                        "q4_uniform", "q5_sampling_from_vector"),
                          hebrew = c("בינומית: שאלה 1", "בינומית: תרגיל", 
                                     "נורמלית: תרגיל 1",
                                     "נורמלית: שאלה 1", 
                                     "נורמלית: תרגיל 2"  ,  "נורמלית: תרגיל 3",
                                     "אחידה: שאלה", "דגימה מוקטור: שאלה"
                                     ))


check_hash_code   = function(hash){
    exercises = exercises_df$exercises
    response_table = learnrhash::decode_obj(hash)
    if (length(response_table)==0){return("Invalid hash code")}
    correct = response_table %>% filter(label %in% exercises, correct) %>% pull(label)
    incorrect = setdiff(exercises, correct)
    if (length(incorrect) == 0){return("Great work! Hash code is valid")}
    else{
      hebrew = exercises_df  %>% filter(exercises %in% incorrect) %>%pull(hebrew)
      print("The following exercises are missing or incorrect:")
      print(hebrew)
    }

} 
}# hash checker
```


## מבוא

במהלך הקורס עד כה התעסקנו רק בניתוח של נתונים אמיתיים. ביחידה זו נלמד
כישור חדש - איך ליצור נתונים על סמך סימולציה (או בעברית - הדמיה) כדי לבדוק רעיונות תיאורטיים.

הרעיון יהיה זהה בכל אחת מהדוגמאות שנפגוש ביחידה זו: 

**1.** אנחנו נניח הנחות מסוימות על המציאות (למשל - שמטבע הוא הוגן)

**2.** נבדוק מה התוצאות שהיינו מצפים לקבל בהינתן הנחות אלו (למשל - את כמות הפעמים שהגיוני שנקבל "עץ" מתוך 10 הטלות אם המטבע היה הוגן)

**3.** נשתמש במידע זה כדי לבחון האם סביר לקבל תוצאה מסויימת בהינתן אותן הנחות (למשל - האם סביר לקבל 10 "עץ" מתוך 10 הטלות)

**4.** ואם המסקנה תהיה שמאוד בלתי סביר שתוצאה כזו תתקבלה - נוכל להסיק שההנחות כנראה לא מתקיימות (והמטבע כנראה לא הוגן)



במהלך היחידה נכיר את השימוש בפונקציות הבאות:

-   `rbinom`
-   `rnorm`
-   `runif`
-   `sample`

המבצעות דגימה אקראית מתוך התפלגות.

נכיר גם הפונקציה `pnorm` המאפשרת לנו לחשב סיכויים על פי התפלגות נורמלית.


## התפלגות בינומית

### דגימה מתוך התפלגות בינומית

הסימולציה הפשוטה ביותר שנוכל לבצע כוללת הדמיה של הטלת מטבע. לפעולה זו יש
2 תוצאות אפשריות - עץ או פלי.

בR אפשר לבצע דגימה כזו באמצעות הפונקציה `rbinom`.

הפונקציה עורכת מספר הטלות מטבע ומחזירה לנו את כמות הפעמים שקיבלנו "עץ".

להלן דוגמה:

```{r coin_flip_example, exercise=TRUE, exercise.eval = FALSE}
# פונקציה שתחזיר לנו 1 או 0, בהסתברות של 0.5
coin_flip = rbinom(n = 1, size = 1, prob = 0.5) 

# נדפיס את התוצאה
print(ifelse(coin_flip==1, "Heads","Tails") )
```

הריצו את הקוד מספר פעמים. כשקיבלנו 1 המטבע נפל על "עץ" וכשקיבנו 0 המטבע
נפל על "פלי".

במידה ונרצה לבצע הטלה של יותר ממטבע אחד נוכל לשנות את הערך של `size`
בפונקציה `rbinom`. כעת הפונקציה תחזיר לנו את מספר הפעמיים שקיבלנו "עץ"
מתוך כל ההטלות שערכנו.

```{r coin_flip_example2,  exercise=TRUE, exercise.eval = FALSE}
# נטיל 100 מטבעות ונספור את כמות הפעמים שקיבלנו עץ
num_heads = rbinom(n = 1, size = 100, prob = 0.5)

print(num_heads)
```

הריצו את הקוד מספר פעמים והתרשמו מטווח הערכים שמתקבלים.

### התפלגות דגימה

במקרה הקודם בדקנו כמה פעמים קיבלו "עץ" מתוך 100 הטלות של מטבע, כשכל הרצה
של הקוד ביטאה חזרה בודדת על הדגימה הזו. כיוון שאנחנו יודעים שהתהליך
המתואר מכיל אלמנט אקראי, נרצה לחזור עליו מספר פעמים כדי לראות מה היא
ההתפלגות של כמויות ה"עץ" שאנחנו מצפים לקבל עבור 100 הטלות מטבע. 

התפלגות זו תתאר לנו את התפלגות הדגימה של כמויות ה"עץ" שמתקבלים מתוך 100 הטלות של מטבע הוגן.

במקרה זה, נוכל ליישם את החזרות המרובות באמצעות שינוי של הערך של `n`
בפונקציה `rbinom`. כשנזין מספר גדול מ1 במקום זה נקבל חזרה וקטור של
תוצאות, כשכל ערך בוקטור מבטא את מספר העצים שהתקבל בסימולציה נפרדת של 100 הטלות מטבע.

כלומר - הארגומנט `n` קובע את מספר החזרות על ה"ניסוי" שלנו, ואילו הארגומנט `size` קובע את מספר ההטלות שכל "ניסוי" כזה כולל. 

```{r coin_flip_sample_dist}
# נבצע 50 חזרות, שכל אחת מהן כוללת 100 הטלות
num_heads = rbinom(n = 50, size = 100, prob = 0.5) 

print(num_heads)
```

קיבלנו וקטור באורך 50, המכילים את התוצאות מ50 חזרות שונות, כשבכל חזרה הטלנו
100 מטבעות.

נחזור על התהליך שוב, אבל הפעם נחזור על הדגימה 10,000 פעם ונציג את
התוצאות בצורה גרפית באמצעות היסטוגרמה.

```{r coin_flip_sample_dist_2}
# נבצע 10,000 חזרות, שכל אחת מהן כוללת 100 הטלות
num_heads = rbinom(n = 10000, size = 100, prob = 0.5)

# נשמור את התוצאות בטבלה לצורך הפקת התרשים
df = data.frame(num_heads = num_heads)

ggplot(df, aes(x = num_heads))+
  geom_histogram( binwidth = 1, fill = "lightblue", color = "black")+
  geom_vline(xintercept = 50 , color = "red", linetype = 2)

```

בגרף ניתן לראות את התוצאות שקיבלנו עבור 10000 חזרות של הטלת מטבע. ניתן
לראות שהתוצאות מתפלגות סביב הערך 50, כלומר כמות ה"עץ" שקיבלנו באופן טיפוסי
עבור 100 ההטלות הייתה סביב 50. ענו על השאלה
הבאה על סמך התרשים:

```{r q1, echo = FALSE}
question("<div style='direction: rtl;'>
מהתרשמותכם, מה הסבירות לקבל פחות מ40 פעמים 'עץ' מתוך 100 הטלות של מטבע הוגן?
</div>",
         
         answer("כמעט וודאי", message = "אמנם היו חזרות בהן קיבלנו פחות מ40 פעמים 'עץ' אבל אילו לא היו מאוד שכיחות"),
         answer("מאוד סביר" , message = "אמנם היו חזרות בהן קיבלנו פחות מ40 פעמים 'עץ' אבל אילו לא היו מאוד שכיחות"),
         answer("לא מאוד סביר, אבל אפשרי",correct = TRUE ),
         answer("בלתי אפשרי", message = "שימו לב, זו תוצאה שיצאה כמה פעמים. היא אמנם לא נפוצה, אבל היא לא בלתי אפשרית" ),
         allow_retry = TRUE,
         
         
  correct="תשובה נכונה. זו תוצאה שאינה שכיחה, אך עדיין אפשרית", 
  incorrect = "נסו שוב"
)
```

למעשה, נוכל לענות על השאלה הנ"ל על ידי חישוב: נעבור על הוקטור שקיבלנו
ונבדוק כמה אחוז מהחזרות הניבו ערך קטן מ40. המספר שנקבל ישקף לנו את
הסיכוי לדגום באופן אקראי פחות מ40 פעמים "עץ" מתוך 100 הטלות מטבע.

```{r p_value_example}
# נבצע 10,000 חזרות, שכל אחת מהן כוללת 100 הטלות
num_heads = rbinom(n = 10000, size = 100, prob = 0.5)

# נחשב את אחוז החזרות בהן קיבלנו פחות מ40 "עץ" מתוך 100 הטלות
p_value = sum(num_heads < 40) / 10000
# הסכימה כאן מחשבת את כמות הפעמים בהם הביטויה בסוגריים נכון (קיבלנו פחות מ40 "עץ")
#  מוהחלוקה ב10000 (מספר החזרות) ממירה את הערך הכמות לשבר עשרוני שמבטא את האחוז המתאים


# ננסח את המסקנה כמשפט מלא
print(paste0("The probability of sampling less than 40 'heads' out of 100 coin flips of a fair coins is approximately ", round(p_value*100),"%"))

# המחשה חזותית
df = data.frame(num_heads = num_heads)

ggplot(df, aes(x = num_heads))+ 
  geom_histogram( bins = 40, aes(fill = num_heads < 40), color = "black")+ # נצבע בצבע אחר את התרשים עבור ערכים הקטנים מ40
  labs(fill = "Less than 40")+ # נשנה את כותרת המקרא
  scale_fill_manual(values = c("lightblue","pink")) # נקבע את הצבעים בהם נעשה שימוש

```

### בירור ערכים קריטיים

אנחנו יכולים גם לבצע את התהליך ההפוך - לבחור הסתברות מסוימת, למשל - 5%, ולבדוק מה בעזרת התפלגות הדגימה את מספר הפעמים שנקבל עץ (או פלי) כך שהסיכוי לקבל כמות כזו או פחות מתוך 100 הטלות של מטבע הוגן הוא 5%.

```{r value_from_p_example}
# 
num_heads = rbinom(n = 10000, size = 100, prob = 0.5)

# פונקציה זו מחזירה לנו את האחוזון ה5 של וקטור התוצאות שלנו
num_heads_5_percent = quantile(num_heads, probs = 0.05) 


print(paste0("There is approximately a 5% chance to get less than ",num_heads_5_percent," heads in 100 coin flips"))

# יצירת תרשים להמחשה
df = data.frame(num_heads = num_heads)
ggplot(df, aes(x = num_heads))+
  geom_histogram( binwidth = 1, aes(fill = num_heads < num_heads_5_percent), color = "black")+
  labs(fill = "Less than 5 percent")+
  scale_fill_manual(values = c("lightblue","pink"))

```

::: text-block1
הפונקציה `quantile` משמשת כדי לחשב אחוזונים מתוך וקטור מספרי. היא מקבלת 2 ארגומנטים: את הוקטור המספרי כארגומנט ראשון, ואת האחוזון שאנחנו מעוניינים בו כארגומנט שני, המכונה `probs`. לדוגמה:

```{r}
quantile(x = 1:100 ,probs = 0.25)
```

הפונקציה הנ"ל מחזירה את הרבעון הראשון (אחוזון 25) של וקטור המכיל את המספרים בין 1 ל100.

ניתן גם לבקש מספר אחוזונים בבת אחת, באמצעות הזנת וקטור בארגומנט  `probs`, למשל:
```{r}
quantile(1:100,prob = c(0.25, 0.5, 0.75))
```
:::


בצורה הזו נוכל לברר את מספר הטלות ה"עץ" שנזדקק לקבל פחות ממנה כדי שהתוצאות שלנו יצביעו על תוצאה מובהקת סטטיסטית. 
  
  
    
`תרגיל`

![](images/Rstudio logo.png){width="100"}

חנה מצאה מטבע ברחוב ורצתה לבדוק אם הוא הוגן (כלומר, שהסיכוי לקבל בו "עץ"
הוא 0.5). היא הטילה את המטבע 50 פעם וקיבלה "עץ" ב20 הטלות. בדקו באמצעות
סימולציה בת 10,000 חזרות מה הסיכוי לקבל את התוצאה שחנה קיבלה או תוצאות קיצוניות יותר (פחות מ20 הטלות "עץ") במידה והמטבע אכן היה הוגן.

פתרו את התרגיל בRstudio וענו על השאלה הבאה:

```{r q2a_hanna, echo = FALSE}

question_text(
  "<div style='direction: rtl;'>מה הסיכוי לקבל 'עץ' 20 פעמים או פחות מתוך 50 הטלות של מטבע הוגן?
  
    
  הזינו את הערך שקיבלתם כמספר בין 0 ל1 (למשל - 0.9 כדי לציין סיכוי של 90%) ועגלו את תשובתכם ל2 מקומות לאחר הנקודה העשרונית
  </div>",
  answer("0.09", correct = TRUE),
  answer("0.1", correct = TRUE),
  answer("0.10", correct = TRUE),
  answer("0.11", correct = TRUE),
  
  allow_retry = TRUE,
  correct="עבודה טובה!", 
  incorrect = "נסו שוב"
)

```


### פתרון אנליטי

למעשה, יכולנו לבחון את השאלות הקודמות גם באמצעות הפתרון האנליטי להתפלגות הבינומית, בעזרת הפונקציה `pbinom`.

פונקציה זו מקבלת מאיתנו תוצאה מסויימת (כמות ה"עץ" שקיבלנו), את מספר ההטלות הכולל ואת הסיכוי הכללי של המטבע לנחות על "עץ" - ומחזירה את הסיכוי לקבל את התוצאה הזו או נמוך ממנה במידה והמטבע אכן היה מאופיין בסיכוי לנחות על "עץ" כפי שהזנו.

החישוב כאן לא נשען על סימולציה - אלא על משוואה מוכרת להתפלגות הבינומית:

$$
P(X = x) = \binom{n}{x} p^x (1 - p)^{n - x}
$$
משוואה זו מתארת את הסיכוי לקבל "עץ" בדיוק x פעמים בהינתן n הטלות של מטבע עם סיכוי p לנחות על "עץ".
כדי לחשב את הסיכוי לקבל כמות מסויימת של "עץ" **או פחות** התוכנה סוכמת את הסיכויים לקבל "עץ" עבור x ועבור כל המספרים הקטנים ממנו.

למשל, אם נרצה לבחון את הסיכוי לקבל 10 "עץ" או פחות מתוך 30 הטלות של מטבע הוגן, נוכל לחשב זאת בצורה הבאה:

```{r pbinom,  exercise=TRUE, exercise.eval = FALSE}
pbinom(q = 10,      # x - כמות ה"עץ" שאנחנו רוצים לחשב את הסיכוי לקבל אותה או פחות ממנה
       size = 30,   # n - כמות ההטלות
       prob = 0.5)  # p - הסיכוי לקבל "עץ" בכל הטלה
```

משמע, בהטלה של 30 מטבעות הוגנים יש סיכוי של 0.05% לקבל 10 "עץ" או פחות.

מאחר ששיטה זו לא נשענת על דגימה עם רכיב של אקראיות - היא תהיה מדויקת יותר מהתוצאות שנקבל בסימולציה.

בסימוציות עם חזרות רבות - ההבדל יהיה זניח, אבל החישוב על סמך המשוואה יאפשר לנו לבצע את החישוב בצורה מהירה יותר. 

## התפלגות נורמלית

רוב המדדים שאנחנו מתעניינים בהם בפסיכולוגיה אינם דיכוטומיים, אלא רציפים.
רבים מהם מתפלגים באופן נורמלי או קרוב לנורמלי באוכלוסייה, ולכן מאוד
שימושי לבצע סימולציות שמערבות דגימה מהתפלגות נורמלית.

הפונקציה שתשמש אותנו במקרים כאלו הינה `rnorm` שמקבלת 3 ערכים: מספר
הדגימות, ממוצע וס"ת.

```{r normal_example}
one_sample = rnorm(n = 1, mean = 8, sd = 2)

many_samples = rnorm(n = 100, mean = 8, sd = 2)

```

נוכל להשתמש בפונקציה זו כדי לבצע הדמיות שעונות על שאלות דומות לאלו
שפגשנו סביב הטלת מטבע, אבל נוגעות למשתנים רציפים. למשל:

מה הסיכוי לדגום באקראי אדם בגובה 140 ס"מ או פחות מתוך אוכלוסיה בה הגבהים
מתפלגים נורמלית עם ממוצע 150 וס"ת 10?

`תרגיל` 
השלימו את הקוד הבא כדי לענות על השאלה הנ"ל.
דגמו 10,000 אנשים מתוך ההתפלגות שצוינה והיעזרו בקוד על מנת לבחון את הסיכוי לדגום אדם בגובה 140 ס"מ או פחות. 

```{r ex1_normal_p_value, exercise=TRUE, exercise.eval = FALSE}
population_distribution = rnorm(n = ____,
                                mean = ____,
                                sd = ____)

height = 140      # הגובה שאנחנו רוצים לברר את הסיכוי לקבל ערך שווה או נמוך ממנו
p_sample_height_or_under = mean(population_distribution<= height)

print(paste0("The chance of sampling a person with a height of ",height," or less is approximately ", round(p_sample_height_or_under *100),"%"))

ggplot(data.frame(height = population_distribution),
       aes(x = height))+ # המחשה חזותית
  geom_histogram( bins = 40, aes(fill = height <= 140), col = "black")+ # נצבע בצבע אחר את התרשים עבור ערכים הקטנים מ40
  labs(fill = "Less than 140")+ # נשנה את כותרת המקרא
  scale_fill_manual(values = c("lightblue","pink")) # נקבע את הצבעים בהם נעשה שימוש

```

```{r ex1_normal_p_value-solution}
population_distribution = rnorm(n = 10000,
                                mean = 150,
                                sd = 10)
```

```{r ex1_normal_p_value-check}

grade_this({
if (abs(user_object_get("p_sample_height_or_under") - 0.15895) < 0.014 ){
     pass("עבודה מצוינת!")
   } else {
     fail("ישנה טעות בקוד. נסו שוב")
   }
})

```


### פתרון אנליטי

כמו בהתפלגות בינומית - גם במקרה הזה ישנה פונקציה המאפשרת לנו לעשות את הבדיקה הזו בצורה
ישירה יותר, מבלי להידרש לדגימה אקראית.

פונקציה זו נקראת `pnorm` והיא בוחנת את הסיכוי לקבל ערך נמוך מערך מסוים
מתוך התפלגות נורמלית עם ממוצע וס"ת נתונים. למשל, את התוצאה מהתרגיל
האחרון יכולנו לחשב כך:

```{r pnorm_example}
height = 140
p  = pnorm(q = height, mean = 150, sd = 10)

print(paste0("The chance of sampling a person with a height of ",height," or less is ", round(p *100),"%"))

```


באופן דומה - ישנה פונקציה שמבצעת את הפעולה ההפוכה. פונקציה זו נקראת `qnorm`, והיא מקבלת מאיתנו סיכוי מסוים, ומחזירה את הערך שהסיכוי לקבל ערך קטן ממנו זהה לסיכוי שהזנו. למשל:

```{r qnorm_example}
p = 0.10
height  = qnorm(p = p , mean = 150, sd = 10)

print(paste0("There is a ", round(p *100),"% chance of sampling a person with a height of ",round(height,1)," or less"))

```

כמו בהתפלגות בינומית, גם כאן החישוב האנליטי יהיה יותר מדויק מאשר התוצאות שנקבל בסימולציה. 


### התפלגות דגימה

נניח שאנחנו לא מתעניינים באדם אחד, אלא בממוצע של קבוצת אנשים.  

במקרים כאלו נוכל להשתמש בפונקציה `rnorm` כדי לדגום מספר ערכים מתוך התפלגות נורמלית מסוימת, המדמה את התפלגות המשתנה באוכלוסיה. מהמדגם שקיבלנו נוכל לחשב ממוצע (או כל סטטיסטי אחר שאנחנו מעוניינם בו). נוכל לחזור על התהליך שוב ושוב באמצעות לולאה, כל פעם עם מדגם אחר, ולחשב כך את התפלגות הדגימה של הממוצע - התפלגות הממוצעים שאנחנו מצפים
לקבל כשאנחנו דוגמים מדגם בגודל מסויים מתוך האוכלוסייה הזו.

נוכל להשתמש בסימולציה כזו כדי לענות על שאלות כמו:

אם היינו דוגמים מדגם של 20 אנשים מתוך אוכלוסייה ששעות השינה בה מתפלגות
נורמלית עם ממוצע 8 וס"ת 2, מה היה הסיכוי לקבל מדגם בודד (גם הוא בן 20
נבדקים) בו הממוצע הוא פחות מ7?

הקוד יראה כך:

```{r norm_sampling_dist_example}
iterations = 10000 # מספר החזרות של ההדמיה

means = c()                # וקטור ריק לטובת שמירת הממוצעים
for (i in 1:iterations){   # קביעת מספר החזרות בהתאם לערך שנקבע
  sample = rnorm(n = 20, mean = 8, sd = 2)  # דגימת 20 ערכים מתוך התפלגות נורמלית עם ממוצע 8 וס"ת 2
  sample_mean = mean(sample)       # חישוב ממוצע המדגם
  means[i] = sample_mean           # הוספת ממוצע המדגם לוקטור התוצאות
}

p_value_less_than_7 = sum(means < 7) / iterations  # חישוב אחוז התצפיות הקטנות מ7

# הדפסת התוצאות כמשפט
print(paste0("The probability of randomly sampling a sample of size 20 with a mean less than 7 from this population is approximately: ", p_value_less_than_7*100," %"))


# הצגת התוצאות כהיסטוגרמה
ggplot(data.frame(means), aes(x = means)) +
  geom_histogram(bins = 30, fill = "lightblue", color = "black")+ # היסטוגרמה
  geom_vline(xintercept = 7, color = "red", linetype = 2)         # קו אנכי מקוקו בערך 7
             
```

ניתוח כזה מאפשר לנו לבדוק עד כמה סביר שמדגם כלשהו הגיע
מאוכלוסיה עם ממוצע וס"ת נתונים. אם נקבל ממוצע שמאוד בלתי סביר לקבל
בדגימה מתוך האוכלוסייה שבדקנו - נוכל להסיק שהמדגם שלנו ככל
הנראה לא הגיע מאוכלוסייה זו.

`תרגיל`

אספתם מדגם של ציונים בשאלון שמחת החיים עבור 30 אנשים. בדקו מה
הסיכויים לקבל מדגם בעל ממוצע הזהה לממוצע שהתקבל ממדגם זה או קטן ממנו -  במידה והייתם דוגמים מתוך אוכלוסייה
של מדינה בה ציוני שמחת החיים מתפלגים נורמלית עם ממוצע 3.5 וס"ת 1.

היעזרו בסימולציה בת 10,000 חזרות.

```{r ex2_sampling_dist_normal, exercise=TRUE, exercise.eval = FALSE, rows=20,exercise.reveal_solution = FALSE}
happiness = c(3.9, 5, 3.5, 2.5, 2.6, 4.2, 5, 2.9, 2.9, 3.8,
              4.1, 2.5, 4.5, 2.5, 2.9, 3, 4.5, 2.6, 2.6, 3.9,
              3.2, 3.5, 3.4, 3.9, 5,4.1, 2.7, 2.3, 3.8, 2.4)





```

```{r ex2_sampling_dist_normal-solution}
happiness = c(3.9, 5, 3.5, 2.5, 2.6, 4.2, 5, 2.9, 2.9, 3.8,
              4.1, 2.5, 4.5, 2.5, 2.9, 3, 4.5, 2.6, 2.6, 3.9,
              3.2, 3.5, 3.4, 3.9, 5,4.1, 2.7, 2.3, 3.8, 2.4)

iterations = 10000 # מספר החזרות של ההדמיה

means = c()       # וקטור ריק לטובת שמירת הממוצעים
for (i in 1:iterations){
  sample = rnorm(n = 30, mean = 3.5, sd = 1)
  sample_mean = mean(sample)
  means = c(means, sample_mean)
}

p_value= sum(means < mean(happiness)) / iterations

print(p_value)

ggplot(data.frame(means), aes(x = means)) +
  geom_histogram(bins = 30, fill = "lightblue", color = "black")+
  geom_vline(xintercept = mean(happiness), color = "red", linetype = 2)

```

```{r q3_sampling_dist_normal, echo = FALSE}

question(
  "<div style='direction: rtl;'>
   בקירוב, מה הסיכוי לדגום מתוך האוכלוסייה של מדינה זו מדגם בגודל 30 בו ממוצע שמחת החיים יהיה זהה או קטן מהממוצע שהתקבל במדגם שלכם?
  </div>",
  answer("40%", correct = TRUE),
  answer("20%"),
  answer("10%"),
  answer("5%"),
  allow_retry = TRUE,
  
           
  correct="כלומר - ממש לא בלתי סביר שהמדגם הגיע מאוכלוסייה זו. יפה מאוד!", 
  incorrect = "נסו שוב"
)

```


#### פתרון אנליטי

גם את השאלות הנוגעות להתפלגות דגימה אנחנו יכולים לפתור ללא
סימולציה - שכן התפלגויות הדגימה של הממוצע של מדגם המתפלג נורמלית מתנהגות
בהתאם לכללים מוגדרים, אותם למדתם בשיעור סטטיסטיקה.

אנחנו יודעים שהתפלגות הדגימה של ממוצע המדגמים שנדמו מאוכלוסייה נורמלית תהיה גם היא נורמלית, עם תוחלת הזהה לתוחלת
ההתפלגות של הערכים באוכלוסייה וטעות תקן (סטיית התקן של התפלגות
הדגימה) השווה לסטיית התקן באוכלוסייה חלקי שורש גודל המדגם:

$$\sigma/\sqrt n$$

להמחשה, לפניכם השוואה בין ההתפלגות שמתקבלת מתוך סימולציה של התפלגות הדגימה אל מול
התפלגות הדגימה התיאורטית שמתקבלת מתוך הכללים הנ"ל:

```{r analytic_solution_example,  exercise = T, exercise.eval = FALSE}

# נתונים על גודל המדגם 
n = 20
# והתפלגות המשתנה באוכלוסייה ממנה אנחנו דוגמים
mean = 8
sd = 2

#### סימולציה ####

iterations = 10000 # מספר החזרות של ההדמיה

means = c()       # וקטור ריק לטובת שמירת הממוצעים
for (i in 1:iterations){      # ביצוע סימולציה של התפלגות דגימה
  sample = rnorm(n = n, mean = mean, sd = sd) # דגימה
  sample_mean = mean(sample)                  # חישוב ממוצע
  means[i] =  sample_mean                     # הוספת הממוצע לוקטור התוצאות
}

#### חישוב אנליטי ####

# התוחלת השל התפלגות הדגימה שווה לתוחלת של האוכלוסייה
# נחשב את טעות התקן (סטיית התקן של התפלגות הדגימה):
se = sd/ sqrt(n)  

#### יצירת תרשים ####

ggplot(data.frame(means), aes(x = means)) +
  geom_histogram(aes(y = ..density..), bins = 30, fill = "lightblue", color = "black",)+   # התפלגות דגימה על סמך סימולציה
  geom_function(fun = dnorm, args = list(mean = mean, sd = se), color = "red", size = 1)+  # התפלגות דגימה - פתרון אנליטי
  annotate(geom = "text", x = 6.7, y = 0.5, label = "Analytic solution", color = "red")+   # הערות בטקסט על גבי התרשים
  annotate(geom = "text", x = 6.7, y = 0.6, label = "Simulation", color = "lightblue")+    # כנ"ל
  theme_minimal()+
  labs(x = "Means", y = "Density")

```

בהתאם, אנחנו יכולים להשתמש בפונקציה `pnorm` גם כדי לבחון את הסיכוי לקבל ממוצע כלשהו מתוך התפלגות הדגימה.

`תרגיל`
חשבו את הסיכוי לאסוף מתוך אוכלוסייה בה ציון רמת הדיכאון מתפלג נורמלית עם תוחלת 20 וס"ת 5 מדגם בן 40 נבדקים עם ציון רמת דיכאון ממוצע שווה או קטן מ18.

```{r analytic_sample_dist_exercise, exercise = TRUE, exercise.eval = FALSE}
p_value = pnorm(q = _____,     # הממוצע שהתקבל במדגם
                mean = ______,   # תוחלת התפלגות הדגימה
                sd = _______)   # סטיית התקן של *התפלגות הדגימה* = טעות התקן 

print(paste("P value =", p_value %>% round(4)))

```

```{r analytic_sample_dist_exercise-check}

grade_this({
  
  if (identical(user_object_get("p_value"),
                pnorm(q = 18,mean = 20, sd = 5/sqrt(40)))
      ) {
    pass("מעולה")
  }
  fail("נסו שוב")
})

```


לחלופין, נוכל להשתמש בפוקנציה `qnorm` כדי לחשב את הערך הקריטי שמתאים לp value מסויים, כך שכל ממוצע שקטן ממנו יהיה מובהק יותר מערך המובהקות שציינו.

`תרגיל` השלימו את התוחלת וטעות התקן מהתרגיל הקודם ומצאו את הערך הקריטי שכל ציון דיכאון נמוך ממנו יצא מובהק סטטטיסטית עבור השערה חד זנבית עם ערך אלפא של 0.05.

```{r analytic_s_dist_exercise2,  exercise = TRUE, exercise.eval = FALSE}
p=0.05
threshold_value = qnorm(p = p,
                        mean = ____,
                        sd = _____)   

print(paste0("For a sample of size 40, any mean depression level smaller than ", threshold_value %>% round(1),
            " would be have a p-value smaller than ", p %>% round(2)))
```

```{r analytic_s_dist_exercise2-check}

grade_this({
  
  if (identical(user_object_get("threshold_value"),
                qnorm(p =user_object_get("p"),
                      mean = 20, 
                      sd = 5/sqrt(40)))
      ) {
    pass("מעולה")
  }
  fail("נסו שוב")
})

```


שתי פונקציות אלו והסימולציות המדמות אותן מאפשרות לנו לבצע בדיקת השערות באצעות מבחן Z למדגם יחיד - נבדוק את הסיכוי לקבל את ממוצע המדגם שלנו מתוך התפלגות דגימה המתאימה להשערת האפס שלנו ואם הסיכוי נמוך מערך רף כלשהו שקבענו (אלפא) נוכל לקבוע שמאוד לא סביר לקבל את התוצאות שלנו אם H0 הייתה נכונה - ולכן אנו מסיקים שהיא ככל הנראה איננה נכונה. 

למשל- בדוגמה הנ"ל, אם היינו דוגמים 40 הנבדקים ומוצאים שממוצע רמת הדיכאוניות שלהם היה 18 - יכולנו להסיק שתוחלת רמת הדיכאוניות באוכלוסייה ממנה דגמנו אינה 20 (בהינתן ההנחה שסטיית התקן של רמת הדיכאוניות הינה 5).


## התפלגויות נוספות
### דגימה מאוכלוסייה שאינה נורמלית

כל הסימולציות שעשינו עד כה עונות על שאלות שיכולנו לפתור גם באופן
אנליטי. הן אמנם עושות זאת באופן ברור יותר ומעט יותר אינטואיטיבי, אבל הן
מפיקות תוצאה זהה למבחנים האנליטיים הקיימים.  

עם זאת, ההדמיות מאפשרות לנו
לבדוק בקלות גם שאלות שאין להן פתרון אנליטי ידוע, או שיש להן פתרון - אבל
הוא מורכב ולא נלמד בקורסי הסטטיסטיקה שלקחנו.


דוגמא אחת שכזו היא מצב בו המדד שאנחנו מתעניינים בו מתפלג באופן שאינו
נורמלי. למשל - כשהוא מתפלג באוכלוסייה באופן אחיד.

כדי לדגום מתוך התפלגויות שאינן נורמליות נוכל להיעזר בפונקציות כדוגמת
`runif` (התפלגות אחידה), `rexp` (התפלגות מעריכית) או פונקציות נוספות
המאפשרות לנו לדגום מתוך התפלגויות אחרות.


#### דגימה מהתפלגות אחידה


  
הפונקציה`runif` משמשת לדגום מתוך התפלגות אחידה. היא מקבלת שלושה ארגומנטים: גודל המדגם, ערך תחתון וערך עליון.
לפניכם דוגמה:

```{r runif_example}
# דוגמה לדגימת 30 תצפיות מהתפלגות אחידה בין הערכים 1 ו2
sample_vec = runif(n = 30, min = 1, max = 2 )

```


`תרגיל` 
![](images/Rstudio logo.png){width="100"}

ידוע שהאורך של שקדי מרק במפעל מסוים מתפלג **באופן אחיד** בין 0.5-1.5 ס"מ.
מה הסיכוי שנדגום באקראי 10 שקדי מרק עם אורך ממוצע של 0.8 ס"מ ומטה?

היעזרו בפונקציה `runif`פתרו את התרגיל בRstudio (באמצעות לולאה עם 10,000 חזרות) וענו על השאלה הבאה:

```{r q4_uniform, echo = FALSE}

question_text(
  "<div style='direction: rtl;'>
  מה הסיכוי שנדגום באקראי 10 שקדי מרק עם אורך ממוצע של 0.8 סנטימטר ומטה?
    
  הזינו את הערך שקיבלתם כמספר בין 0 ל1 (למשל - 0.9 כדי לציין סיכוי של 90%) ועגלו את תשובתכם ל2 מקומות לאחר הנקודה העשרונית
  </div>",
  answer("0.01", correct = TRUE),
  answer("0.02", correct = TRUE),
  allow_retry = TRUE,
  
           
  correct="מעולה!", 
  incorrect = "נסו שוב"
)

```


### דגימה מתוך וקטור

סוג התפלגות הדגימה האחרון שנשתמש בו בפרק זה נוגע לדגימה מתוך וקטור. כאן
אנחנו נספק את הערכים האפשריים ונבצע מתוכם דגימה באקראי. לצורך כך נשתמש
בפונקציה `sample`.

לדוגמה, אם נרצה לדגום באקראי 3 הטלות קובייה נוכל לעשות זאת בצורה הבאה:

```{r sample_example}
three_rolls = sample(x= 1:6,      # וקטור המכיל את המספרים בין 1 ל6, ממנו תתבצע הדגימה
                     size = 3,    # כמות הדגימות שנרצה לדגום
                     prob = rep(x = 1/6,times = 6), # הסיכוי לדגום כל אחד מהערכים בוקטור
                     # במקרה הזה לא היינו חייבים לציין את הסיכויים - שכן ברירת המחדל
                     # היא סיכויים שווים לכל ערך
                     replace = T) # משתנה השולט בהאם הדגימה תהיה עם או בלי החזרה
                                  # במידה ונדגום ללא החזרה - לא נוכל לקבל פעמיים את
print(three_rolls)                # אותה התוצאה, ולכן ביקשנו כאן דגימה עם החזרה
```

`תרגיל`

![](images/Rstudio logo.png){width="100"}

בנו התפלגות דגימה (על סמך 10,000 חזרות) עבור הערך המקסימלי מתוך הטלה של ארבעה קוביות.

היעזרו בפונציה `max` כדי לחלץ מתוך וקטור את הערך המקסימלי.


```{r q5_sampling_from_vector, echo = F}
question_text(
  "<div style='direction: rtl;'>
  מה הסיכוי שהערך המקסימלי מתוך ארבעה ההטלות יהיה קטן מ3?
  
  הזינו את הערך שקיבלתם כמספר בין 0 ל1 (למשל - 0.9 כדי לציין סיכוי של 90%) ועגלו את תשובתכם ל2 מקומות לאחר הנקודה העשרונית
  </div>",
  answer("0.01", correct = TRUE),
  answer("0.02", correct = TRUE),
  allow_retry = TRUE,
  
           
  correct="מעולה!", 
  incorrect = "נסו שוב"
)


```

##  מבחני עוצמה (הרחבת רשות)

שימוש נוסף, ואחרון עבור יחידה זו, לסימולציות משמש אותנו לחישוב הסיכוי לקבל תוצאה מובהקת בניסוי.

נניח שאנחנו רוצים לבצע ניסוי כלשהו - אבל לא בטוחים אם גודל המדגם שלנו מספיק גדול כדי שנצליח לקבל תוצאה מובהקת. כדי לבדוק זאת נוכל לבצע סימולציה בה אנחנו:

**1)** מניחים שהאפקט שאנחנו מחפשים אכן קיים בעולם  
**2)** דוגמים מתוך התפלגות שמתאימה לאפקט זה מדגם בגודל הזהה למדגם שאנחנו מתכננים לאסוף בניסוי האמיתי שלנו   
**3)** בודקים אם התוצאה שקיבלנו במדגם המסומלץ מספיקה כדי לדחות את השערת האפס שלנו  
**4)** חוזרים על התהליך מספר רב של פעמים, ובודקים בכמה אחוז מהמקרים קיבלנו תוצאות מובהקות סטטיסטית

תוצאה של ניתוח כזה תוכל לעזור לנו להעריך את הסיכויים שלנו לקבל תוצאה מובהקת בניסוי שלנו, תחת ההנחה שהאפקט שאנחנו מחפשים קיים. 

מאפיין זה של הניסוי המתוכנן מכונה "עוצמת המבחן".

  
### עריכת מבחן בודד
  
לפני שניגש לסימולציה של עוצמת מבחן נתחיל בלערוך מבחן בודד, כפי שעשינו מוקדם יותר ביחידה באמצעות הפונקציה `pnorm`.


המבחן שנערוך, הנקרא מבחן Z לתוחלת מדגם בודד, בוחן את הסיכוי לקבל מדגם בעל ממוצע מסוים במידה ודגמנו מאוכלוסייה עם ממוצע וס"ת מוגדרים - ואם הסיכוי לקבל תוצאה כזו קטן מדי (לרוב - קטן מ0.05) נסיק שהקבוצה שדגמנו ממנה כנראה אינה שייכת לאוכלוסייה זו.


למשל: נבחן האם אנשים שלומדים פסיכולוגיה נמוכים יותר בתכונת הנוירוטיות ביחס
לאוכלוסיית הסטודנטים הכללית, בהינתן וידוע שבקרב כלל הסטודנטים תכונה זו
מתפלגת נורמלית עם ממוצע 3 וס"ת 1.

נערוך את המבחן באמצעות הפונקציה `pnorm`, עבור המקרה בו אספנו מדגם של 30 סטודנטים וקיבלנו עבורם ציון נוירוטיות ממוצע של 2.5:

```{r z_test_example}
# נניח ואספנו מדגם של 30 סטודנטים וקיבלנו עבורם ציון נוירוטיות ממוצע של 2.5
mean_score = 2.5
sample_size = 30

p_value = pnorm(mean_score ,                # הממוצע שקיבלנו
                mean = 3,                   # הממוצע תחת השערת האפס
                sd = 1/sqrt(sample_size))   # טעות התקן
print(p_value)
```

ערך הp value שקיבלנו קטן מ0.05 ולכן נוכל לדחות את השערת האפס. כלומר - אם נקבל ממוצע נוירורטיות של 2.5 ממדגם של 30 סטודנטים לפסיכולוגיה, נוכל להסיק שרמת הנוירוטיות של סטודנטים לפסיכולוגיה קטנה מ3, בהינתן שסטיית התקן של רמת הנוירוטיות באוכלוסייה הינה 1.

### דגימה מתוך ההשערה האלטרנטיבית

בפועל, גם אם המצב האמיתי בעולם הוא שהתוחלת של הנוירוטיות של סטודנטים לפסיכולוגיה היא באמת 2.5 - עדיין מאוד סביר שבדגימה אקראית כלשהי נקבל מדגם עם ממוצע גבוה יותר , כך שייתכן שנקבל מדגם שלא יאפשר לנו לדחות את השערת האפס.

כדי להבין את הסיכויים לכך נערוך סימולציה: 

נדמה את התהליך בו אנחנו דוגמים מדגם של 30 ערכים מתוך אוכלוסייה שכזו (המקיימת את ההשערה האלטרנטיבית, ולכן תוחלת הנוירוטיות בה שווה ל2.5), ונבחן האם המדגם שלנו אפשר לנו לדחות את השערת האפס.
 
נחזור על התהליך מספר רב של פעמים ונבחן בכמה אחוז מהמקרים המדגם שקיבלנו היה עם ממוצע נמוך מספיק.

התוצאה שנקבל תאפשר לנו להעריך את הסיכוי שלנו לדחות את השערת האפס על סמך מדגם בגודל 30, במידה והמצב האמיתי בעולם הוא שהתוחלת הינה 2.5.

```{r power_example}
sample_size = 30   # גודל המדגם

p_values = c()     # וקטור ריק לשמירצת התוצאות

for (i in 1:1000){
  sampled_neuroticism = rnorm(sample_size, mean = 2.5, sd = 1)   # נדגום 30 ערכים
  sample_mean = mean(sampled_neuroticism )                       # נחשב את הממוצע של המדגם
  
  p_value = pnorm(sample_mean,      # ביצוע המבחן עבור המדגם שקיבלנו
                  mean = 3,        
                  sd = 1/sqrt(sample_size))
  p_values[i] =  p_value          # הוספת ערך המובהקות לוקטור התוצאות
}

percent_significant = mean(p_values < 0.05) # חישוב אחוז המקרים בהם קיבלנו תוצאה מובהקת

print(paste0("Power = ",round(percent_significant,2))) # דיווח


```

כלומר - אם המצב בעולם הוא שתכונת הנוירוטיות אצל סטודנטים לפסיכולוגיה אכן
מתפלגת נורמלית עם ממוצע 2.5 וס"ת 1, אז
בכ`r round(percent_significant*100)` אחוז מהמקרים מדגם בגודל 30 יניב
תוצאות מובהקות סטטיסטית ויראה שהם אכן נמוכים יותר בתכונה זו מאוכלוסיית
הסטודנטים הכללית.

### תרגיל רשות

![](images/Rstudio logo.png){width="100"}

`תרגיל` בצעו ניתוח עוצמה למקרה הבא
  
חוקר מעוניין לבחון האם החתולים בקמפוס חמודים יותר מחתולים בשאר העיר.
מכיוון שידוע שציוני החמידות של חתולים בירושלים מתפלגים נורמלית עם תוחלת
6 וס"ת 1.5, הוא מתכנן לאסוף מדגם של חתולים בקמפוס ולהשוות את הממוצע שלהם
להתפלגות הדגימה של אוכלוסייה זו.
  
בהינתן וציוני החמידות של החתולים **בקמפוס** מתפלגים נורמלית עם תוחלת 6.5
וס"ת 1.5 - מה הסיכוי שמדגם של 40 חתולים יספיק לחוקר כדי לקבל תוצאות מובהקות
סטטיסטית?\

הערה: שימו לב שההשוואה שעורך החוקר בודקת אם החתולים בקמפוס חמודים
**יותר** אילו ערך הp value שמתקבל מהפונקציה `pnorm` בוחן את הסיכויים
לקבל ערך **נמוך** יותר - ולכן צריך לחשב את הערך המשלים ל1 לערך שתקבלו: `p_value = 1-pnorm(...)`.

פתרו את התרגיל בRstudio וענו על השאלה הבאה:


```{r q6_power, echo = FALSE}

question_text(
  "<div style='direction: rtl;'>
  מה תהיה עוצמת המבחן עבור ניסוי שכולל מדגם בגודל 40?
  
  הזינו את הערך שקיבלתם כמספר בין 0 ל1 (למשל - 0.9 כדי לציין עוצמה של 90%) ועגלו את תשובתכם ל2 מקומות לאחר הנקודה העשרונית
  </div>",
  answer("0.7", correct = TRUE),
  answer("0.70", correct = TRUE),
  answer("0.69", correct = TRUE),
  answer("0.68", correct = TRUE),
  answer("0.67", correct = TRUE),
  allow_retry = TRUE,
  
  correct="מעולה!", 
  incorrect = "נסו שוב"
)

```

התרגיל הינו תרגיל רשות, כלומר - ניתן להגיש את היחידה גם מבלי לענות על השאלה הנ"ל.


## משוב

בחלק זה נבקש את המשוב שלכם על הלומדה. אנא ענו בכנות וביסודיות, על מנת שנוכל להשתפר. התשובות ישמרו בצורה אנונימית ולא ישפיעו על בדיקת התרגיל עצמה.

שימו לב: ברוב השאלות אין אפשרות לשנות את התגובה לאחר לחיצה על כפתור ההגשה.

```{r survey_q1, echo = FALSE}
question(#"כמה קשה הייתה היחידה עבורך?",
          "מה הייתה רמת הקושי של הלומדה עבורך?",
         type="single",
         answer("קלה"        ,correct=TRUE,  message = "תגובתך: קלה"),
         answer("בינונית"    ,correct=TRUE,  message = "תגובתך: בינונית"),
         answer("מאתגרת"     ,correct=TRUE,  message = "תגובתך: מאתגרת"),
         answer("קשה"        ,correct=TRUE,  message = "תגובתך: קשה"),
         answer("קשה מאוד"   ,correct=TRUE,  message = "תגובתך: קשה מאוד"),
         correct = "",
         incorrect = "",
         allow_retry = TRUE
)
```

```{r survey_q2, echo = FALSE}
question(#"עד כמה את/ה מרגיש/ה שהבנת את החומר שנלמד ביחידה?",
        "עד כמה החומר הנלמד ביחידה הועבר בצורה ברורה לדעתך?",
         type="single",
         answer("במידה רבה"      , correct=TRUE,  message = "תגובתך: במידה רבה"),
         answer("במידה סבירה" ,    correct=TRUE,  message = "תגובתך: במידה סבירה"),
         answer("באופן חלקי"  ,    correct=TRUE,  message = "תגובתך: באופן חלקי"),
         answer("במידה מעטה"  ,correct=TRUE,  message = "תגובתך: במידה מעטה"),
         answer("במידה מעטה מאוד",correct=TRUE,  message = "תגובתך: במידה מעטה מאוד"),
         correct = "",
         incorrect = "",
         allow_retry = TRUE
)
```

```{r survey_q3, echo = FALSE}
question("כמה זמן לקח לך לפתור את הלומדה בערך?",
         type="single",
         answer("פחות משעה"      , correct=TRUE,  message = "תגובתך: פחות משעה"),
         answer("שעה" ,    correct=TRUE,  message = "תגובתך: שעה"),
         answer("שעתיים"  ,    correct=TRUE,  message = "תגובתך: שעתיים "),
         answer("שלוש שעות"  ,correct=TRUE,  message = "תגובתך:  שלוש שעות"),
         answer("ארבע שעות ומעלה",correct=TRUE,  message = "תגובתך:  ארבע שעות ומעלה "),
         correct = "",
         incorrect = "",
         allow_retry = TRUE)
```

```{r survey_q4, echo = FALSE}
question("אילו נושאים בלומדה היו קשים במיוחד עבורך?",type="learnr_checkbox",
         answer("התפלגות בינומית", correct=TRUE),
         answer("התפלגות נורמלית", correct=TRUE),
         answer("התפלגות אחידה", correct=TRUE),
         answer("דגימה מוקטור", correct=TRUE),
         answer("התפלגות דגימה", correct=TRUE),
         answer("מבחני עוצמה", correct=TRUE),
         answer("אף נושא", correct=TRUE),
         
         correct = "תגובתך התקבלה",
         try_again = "תגובתך התקבלה",
         incorrect = "תגובתך התקבלה",
         allow_retry = TRUE
)
```

```{r survey_q5, echo = FALSE}
  question_text(    "מצאת טעות? ספר/י לנו עליה (יש להקפיד ולציין את שם הפרק)",
   # "הערות ספציפיות לגבי חלקי היחידה השונים",
    incorrect = "תגובתך התקבלה",
    correct = "תגובתך התקבלה",
    answer(text = "/'ק/'ק",correct = TRUE),
    rows = 10,
    trim = FALSE
  )
```

```{r survey_q6, echo = FALSE}
  question_text(
        "יש לך עוד משהו לספר לנו? נשמח לשמוע משוב מפורט, הצעות כלליות או ספציפיות לפרק מסוים לגבי יחידה זאת.",
   # "הערות כלליות לגבי היחידה",
    incorrect = "תגובתך התקבלה",
    correct = "תגובתך התקבלה",
    answer(text = "/'ק/'ק",correct = TRUE),
    rows = 10,
    trim = FALSE
  )
```


## הגשה

עברו על הקובץ וודאו שהגשתם את כל התרגילים ועניתם על כל השאלות

במידה וכל התשובות שלכם תקינות יש ללחוץ על הכפתור: Generate, להעתיק את
הטקסט שמופיע בחלון למטה ולהגישו במודל\
בהצלחה!

```{r context="server"}
learnrhash::encoder_logic()
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui()
```


הדביקו את קוד הhash בתוך הגרשיים בפונקצייה הבאה כדי לוודא שעניתם על כל השאלות והתרגילים בלומדה.

```{r hash_check,  exercise = TRUE, exercise.eval = FALSE}

check_hash_code("")

```


