---
title: "Lesson 02: Variables"
author: "Copyright 2024 Psychology Department. Hebrew University of Jerusalem. All rights reserved"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: false
description: >
  <div style='direction: rtl;'>
  בשיעור זה נלמד על אבני הבניין הבסיסיים של שפת התכנות - האופרטורים (פעולות) והמשתנים.
  </div>
runtime: shiny_prerendered
---

<style>
h1, h2, h3, h4, h5, h6 {
direction: rtl;
}
p {
direction: rtl;
}
</style>


<!-- tutorial options and checking options - TODO set params -->
```{r setup, include=FALSE}
library(learnr)
library(gradethis)
##### פונקציה לבדיקה
exercises_df = data.frame(exercises = c("hello_world_ex", "basic_op_example_a", "ari_op_ex", "ari_op_ex2", "ari_op_ex3",
  "comp_op_ex", "comp_op_ex2", "comp_op_ex3",
  "logi_op_ex1", "logi_op_ex2", "logi_op_ex3",
  "var_example_a", "string_ex1", "str_comp_ex", "log_ex",
  "final_ex1", "final_ex2", "final_ex3", "intro_question","var_question", "numeric_question", "str_question"), hebrew = c("שלום עולם",
  "תרגיל פעולות בסיסיות",
  "תרגיל פעולות אריתמטיות 1",
  "תרגיל פעולות אריתמטיות 2",
  "תרגיל פעולות אריתמטיות 3",

  "תרגיל השוואה 1",
  "תרגיל השוואה 2",
  "תרגיל השוואה 3",

  "תרגיל לוגי 1",
  "תרגיל לוגי 2",
  "תרגיל לוגי 3",

  "תרגיל משתנים ",
  "תרגיל מחרוזת ",
  "תרגיל השוואת מחרוזת ",
  "תרגיל השוואה לוגית",

  "תרגיל מסכם 1",
  "תרגיל מסכם 2",
  "תרגיל מסכם 3", 
  "שאלת פתיחה", 
  "שאלת משתנים", 
  "שאלת מספרים", 
  "שאלת מחרוזות"))


check_hash_code   = function(hash){
    exercises = exercises_df$exercises
    response_table = learnrhash::decode_obj(hash)
    if (length(response_table)==0){return("Invalid hash code")}
    correct = response_table %>% filter(label %in% exercises, correct) %>% pull(label)
    incorrect = setdiff(exercises, correct)
    if (length(incorrect) == 0){return("Great work! Hash code is valid")}
    else{
      hebrew = exercises_df  %>% filter(exercises %in% incorrect) %>%pull(hebrew)
      print("The following exercises are missing or incorrect:")
      print(hebrew)
    }

} 

knitr::opts_chunk$set(error = TRUE)
gradethis::gradethis_setup()
```

##  היכרות עם הלומדה

ברוכות וברוכים הבאים!

קודם כל נכיר קצת את הלומדה איתה נעבוד בשנה הקרובה.  
בלומדה נוכל להראות לכם קטעי קוד אותם לא תוכלו לערוך, למשל:  

```{r}
print("hello world")
```
קטע קוד זה פשוט מדפיס את ההודעה "hello world". אתם יכולים לראות את הקוד והתוצאה אבל לא לערוך אותה. כשתראו קטעי קוד כאלה המטרה שלהם היא ללמד אתכם את כתיבת הקוד. אם תרצו לראות איך שינויים שונים ישפיעו תמיד תוכלו להעתיק את הקוד לrstudio, לערוך ולהריץ אותו.  

כאשר תקבלו מאיתנו תרגיל תוצג לכם תיבת קוד, לפעמים היא תהיה ריקה ואז עליכם לכתוב את הקוד מההתחלה, לפעמים חלק מהקוד יהיה קיים בתיבה ותדרשו לערוך אותו.  
למשל בתיבת הקוד הבאה כתבו את הקוד שידפיס את ההודעה "hello world" בעצמכם. 

```{r hello_world_ex, exercise=TRUE, exercise.eval = FALSE}

```

```{r hello_world_ex-solution}
print("hello world")
```

```{r hello_world_ex-check}

grade_this_code(correct="איזה יופי! צלחתם את התרגיל הראשון", incorrect = "עוד לא הגעתם לתשובה הנכונה, צריך להכניס את הקוד המופיע למעלה בדיוק.")
```

שימו לב ל4 הכפתורים שמוצגים לכם:  
`start over` - יאפס את התיבה (כלומר ימחק את כל מה ששינתם) לכן ללחוץ בזהירות.  
`solution` - פתרון התרגיל. לא תמיד האופציה תהיה זמינה. אם היא איננה זמינה סימן שאתם צריכים לפתור את התרגיל לבד.  
`run code` - הרצת הקוד, לאחר לחיצה על אופציה זאת תראו את התוצאה של הקוד שלכם.   
`submit answer` - הגשת התרגיל, שימו לב, הרצת הקוד לא מגישה את התרגיל, חייבים ללחוץ על כפתור זה כדי שהתרגיל יוגש ויבדק. ניתן להגיש כמה פעמים וההגשה הסופית היא זאת שתיבדק. לאחר לחיצה על כפתור זה יופיע לכם משוב, על מנת לקבל ציון עובר בתרגיל יש לפתור נכונה את הרוב המוחלט של התרגילים.  

לאחר שהבנתם, הגישו את התרגיל בתיבה מעלה.

לפעמים גם יופיעו לכם שאלות בחירה מרובה או שאלות פתוחות. עליכם לענות על כל השאלות ולהגיש גם אותן. למשל:  

```{r intro_question, echo = FALSE}
question(" מוכנים להתחיל?",
         answer("כן! אני רוצה לכתוב קוד! ", correct = TRUE),
         answer("עוד לא..", message = "אז חזרו כשאתם מוכנים"),
         allow_retry = TRUE
)
```


##  פעולות בסיסיות - מבוא

ועכשיו נעבור לנושא הראשון שלנו ללומדה זאת - פעולות בסיסיות (או Operators באנגלית).  
באופן הבסיסי ביותר, אנחנו יכולים להשתמש בתיבת הקוד כמו במחשבון. בתיבת הקוד מוזן כרגע תרגיל פשוט של חיבור - 1+1.  
 כאשר תלחצו על הכפתור המתאים ותריצו את הקוד תודפס לכם התוצאה. נסו לערוך את הקוד ולחבר מספרים שונים. כדי לראות את התוצאה מתעדכנת יש להריץ את שורת הקוד כל פעם מחדש.  
 לאחר שניסיתם מספר תרגילים צרו תרגיל שהתוצאה שלו תהיה 5 והגישו את הקוד

```{r basic_op_example_a, exercise=TRUE, exercise.eval = FALSE}
1 + 1
```

```{r basic_op_example_a-check}
grade_this({if (identical(.result, 5)) {
    pass("מצויין!.")
}
  fail("התוצא עוד לא 5, נסו שוב!")
})
```

כנראה שכרגע עוד לא ברור לכם איך חיבור מספרים באמצעות קוד רלוונטי לפסיכולוגיה. אל דאגה, כבר ביחידה זו נראה איך התכנות יכול לעזור לנו במחקר.

בקורס זה נעבוד עם נתונים אמיתיים שנוכל להתחיל להכיר כבר עכשיו.

`היכרות עם הנתונים`  
בקורס זה נשתמש בנתונים ממחקר אמיתי שנערך על סטודנטים. כחלק ממחקר זה הסטודנטים ענו על שאלונים בנושא אישיות בתרגיל זה נשתמש בנתוני השאלונים שלהם.  
השאלון מורכב מ10 שאלות אשר בוחנות 5 רכיבים אישיותיים.

נועה (שם בדוי) היא סטודנטית לפסיכולוגיה בשנה השלישית שענתה על השאלון ועם נתוניה נעבוד היום.

זהו שאלון אמיתי לחלוטין ותלמדו עליו עוד בקורס "אישיות" בהמשך התואר. רוצים לדעת עוד כבר עכשיו? לחצו על הקישור:
[Big5](https://he.wikipedia.org/wiki/%D7%97%D7%9E%D7%A9_%D7%94%D7%AA%D7%9B%D7%95%D7%A0%D7%95%D7%AA_%D7%94%D7%92%D7%93%D7%95%D7%9C%D7%95%D7%AA)  
* אם הקישור לא נפתח לחצו על עליו בליווי לחיצה על מקש control

שאלון זה מודד 5 תכונות אישיות, בנתונים אלה כל תכונת אישיות נמדדת באמצעות 2 שאלות.  
במאגר הנתונים הללו יש 2 סוגים של שאלות, שאלות רגילות ושאלות הפוכות.  
למשל אם נרצה למדוד את תכונת ה"מצפוניות" (באנגלית Conscientiousness). תכונה אשר מתארת עד כמה האדם מאופיין במוסר עבודה. דוגמה לשאלה רגילה תהיה:  
"אני אדם שחשוב לו לסיים משימות שהתחיל"  
מי שעונה על השאלון ידרג עד כמה היגד זה נכון בין 1-7, כך שככל שהדירוג גבוה יותר כך האדם יותר מאופיין במצפוניות.  
דוגמה לשאלה הפוכה היא:  
"אני משאיר דברים לרגע האחרון"  
כך שמי שמדרג את היגד זה גבוה דווקא מאופיין ברמה נמוכה של מצפוניות.  

לאחר שהכרנו את השאלון נוכל להתחיל לחשב מדדים שונים. 


## פעולות אריתמתיות
פעולות אריתמתיות (Arithmetic operators) הסוג הראשון והפשוט ביותר הן פעולות החשבון הבסיסיות שכולנו מכירים, זהו מקרא של כל הפעולות הבסיסיות:  
`+` - חיבור  
`-` - חיסור  
`*` - כפל  
`/` - חילוק  
`^` - חזקה  

`דוגמה` 
נתחיל לנתח את תשובותיה של נועה ונתחיל עם תכונת האישיות "מצפוניות" בקיצור נקרא לה C.
תשובותיה של נועה לשתי השאלות הבוחנות מצפוניות הן 6 ו2. כעת נחשב את הציון הכולל:

```{r}
6+2
```

כדי לקבל את הציון הממוצע לתכונה C עלינו לחלק את הציון הכולל במספר השאלות, כך:

```{r}
# עכשיו נצטרך להשתמש בשתי פעולות!
# שימו לב לשימוש בסוגריים!
# התוכנה פועלת לפי סדר פעולות חשבון (קודם חילוק וכפל ואז חיבור וחיסור) אם נרצה שפעולת החיבור תתבצע קודם נשים אותה בסוגריים
(6+2)/2
```

שמתם לב שיש קצת פער בין התשובות? זה בעצם בגלל שהשאלה השניה היא "שאלה הפוכה" כלומר היא ביקשה מנועה לדרג תכונה שקשורה למצפוניות בצורה הפוכה (למשל עד כמה היא נוטה לאחר בהגשת תרגילים), כך שציון נמוך דווקא מעיד על מצפוניות גבוהה. 

בואו נחשב את הציון הממוצע עכשיו עם תיקון של השאלה ההפוכה.  
תיקון של השאלה ההפוכה נבצע באמצעות החסרה של הציון ההפוך מציון המקסימלי בשאלון +1.  

הציון הגבוה ביותר בשאלון הוא 7 ולכן נרצה להחסיר את התשובה ההפוכה מ8. הסיבה שזאת הפעולה הנכונה היא שהיא מאפשרת להפוך את הסדר היחסי. למשל אם מישהו דירג 1 בשאלה הפוכה (הציון הכי נמוך) נרצה שלאחר ההיפוך יהיה לו את הציון הגבוה ביותר (7).  
במקרה שלנו כדי לחשב את הציון הממוצע נבצע את החישוב הבא:

```{r}
(6 + 8-2)/2
```

`ועכשיו נתרגל! `  
נתונות לכם שתי תשובות נוספות של נועה על תכונת האישיות "מוחצנות " (בקיצור E). שאלה רגילה בה נועה דירגה 3 ושאלה הפוכה בה דירגה 6. 

חשבו את הציון הממוצע של נועה בתכונה זאת (שימו לב לשאלה ההפוכה)

```{r ari_op_ex, exercise=TRUE, exercise.eval = FALSE}

```

```{r ari_op_ex-solution}
(3 + 8-6)/2
```

```{r ari_op_ex-check}
grade_this({if (identical(.result, (3 + 8-6)/2)) {
    pass("מצויין! חישבתם  נכון.")
}
  fail("זאת לא התשובה הנכונה.. נסו שוב!")
})
```


הנה כמה נתונים כלליים על נועה:
גיל - 24  
שנות לימוד - 13  
עזרו לנו לחשב מדד יחסי לכמות הזמן במהלך חייה שנועה בילתה במסגרת חינוך. חשבו זאת על ידי הכפלת שנות הלימוד ב2 וחלוקת התוצאה בגיל.  

```{r ari_op_ex2, exercise=TRUE, exercise.eval = FALSE}

```

```{r ari_op_ex2-solution}
13*2/24
```

```{r ari_op_ex2-check}
grade_this({if (identical(.result, 13*2/24)) {
    pass("נכון מאוד!")
}
  fail("זאת לא התשובה הנכונה.. נסו שוב!")
})
```

נודע לנו שבגרסה קודמת של הניסוי נועה ענתה על תכונת "מצפוניות" 
בתחילה היא דירגה שאלה רגילה בציון 7, ושאלה הפוכה בציון 3.

כעבור שבועיים היא ענתה שוב על השאלון, וציוניה השתנו: שאלה רגילה 6, שאלה הפוכה 2.

חשבו את השינוי (ההפרש) בציון הממוצע של נועה בתכונה זו בין הפעם הראשונה לשנייה (לאחר תיקון השאלה ההפוכה בשני המקרים).
```{r ari_op_ex3, exercise=TRUE, exercise.eval = FALSE}

```

```{r ari_op_ex3-solution}
((7 + (8-3))/2) - ((6 + (8-2))/2)
```

```{r ari_op_ex3-check}
grade_this({if (identical(.result, ((7 + (8-3))/2) - ((6 + (8-2))/2))) {
    pass("מעולה! הגעתם לתשובה")
}
  fail("עוד לא הגעתם לתשובה!")
})
```


## פעולות יחסיות
פעולות יחסיות (Relative operators) הן פעולות שנועדו לאפשר לנו להשוות בין שני משתנים שונים.
בניגות לפעולות אריתמתיות כאשר נשתמש בפעולות מסוג זה התוצאה תהיה "נכון" או "לא נכון" (TRUE או FALSE) ולא מספר. שכן התוצאה יכולה להיות שההשוואה שפירטנו מתקיימת או שלא. להלן מקרא של כל הפעולות הלוגיות:  

`>` - גדול מ  
`<` - קטן מ  
`>=` - גדול שווה ל  
`<=` - קטן שווה ל  
`==` - שווה ל  
`!=` - שונה מ  

נתחיל בהשוואה מאוד פשוטה:
```{r}
3 < 4
```

במקרה הזה שאלנו האם 4 גדול מ3? התשובה כמובן חיובית (TRUE).
במקרה זה לא היינו צריכים את הקוד, התשובה הייתה ברורה לנו ואולי קשה לכם להבין למה נצטרך לבצע השוואה כזאת. אבל, בשפת התכנות אנחנו עשויים לעבוד עם חישובים מורכבים בהרבה, למשל:

```{r}
(2395*25)/16 > 55^9
```
בדוגמה השניה כבר ראינו שני חישובים מאוד מורכבים. אבל אם התוצאה של החישובים לא משנה לנו אנחנו יכולים להבין מהקוד שהתוצאה בצד שמאל של המשוואה לא גדולה מהתוצאה בצד ימין.

בנוסף להשוואה בין שני ערכים וקביעה האם אחד גדול מהשני אפשר גם לבדוק האם שני ערכים זהים זה לזה.  

`שימו לב!` ערך השוואה לגבי האם הערכים זהים נעשית עם סמל כפול בכוונה, לסימן = יחיד יש משמעות שונה לחלוטין.  

למשל אם נחזור לדוגמה של נועה נוכל לחשב האם הציון הממוצע שלה בC זהה לציון הממוצע שלה בE למשל:

```{r}
((6 + 8-2)/2) == ((3 + 8-6)/2)

```


   לאחר שהבנו שלנועה יש מידת מצפוניות שונה ממידת המוחצנות שלה. נרצה לדעת האם היא יותר מצפונית ממוחצנת?


```{r}
#נבדוק האם ציון מצפוניות הממוצע גדול מציון המוחצנות
((6 + 8-2)/2) > ((3 + 8-6)/2)
```
והתשובה היא כן!

  
עכשיו גם נתון לנו שהציון הממוצע שלה בתכונה נעימות (A) היא 6.
כעת נרצה לבדוק האם היא נחמדה לפחות כמו שהיא מצפונית?


```{r}
# במחשבה ראשונה ניתן להשוות כך.
((6 + 8-2)/2) < 6
```

קיבלנו תשובה שלילית, אך בעצם ההשווא שביצענו כאן לא ממש עונה על השאלה של "לפחות כמו" אלא על השאלה "גדול מ". 
כשנישאל האם מספר אחד גדול לפחות כמו השני והמספרים שווים נרצה לקבל תשובה חיובית ולכן נשתמש באופרטור אחר:

```{r}
((6 + 8-2)/2) <= 6
```

ועכשיו התשובה חיובית. כצפוי.

למדנו להכיר אופרטור שבודק האם שני ערכים הם זהים ואופרטורים שמשווים את הגודל של שני ערכים. האופרטור האחרון (not) בודק האם שני ערכים שונים זה מזה.  
כלומר, הוא יחזיר שההשוואה נכונה כל עוד שני הערכים לא זהים. למשל:

```{r}
5 != 7

```

כשמשווים בין שני מספרים פשוטים קצת קשה להבין איך האופרטור הזה יכול לשמש אותנו. אבל כאשר נתחיל לעבוד עם מאגרי נתונים גדולים האופציה לבדוק ששני ערכים שונים היא שימושית לא פחות.
למשל במקרה ושי לנו בנתונים שני נבדקים שענו על השאלון בצורה זהה ונצטרך להבין האם אלה תוצאות אמיתיות, כלומר, שני משתתפים שונים  עם אותן תשובות או שורה שהועתקה פעמיים נוכל להשוות את מספרי הזהות שלהם ולוודא שהם שונים.

`תרגיל`  

קודם כל תזכורת על ציוניה של נועה שישמשו אתכם לפתרון של השאלות:  
 
ציוני המצפוניות:  
6   
2 (שאלה הפוכה)  
ציוני המוחצנות:  
3  
6 (שאלה הפוכה)  
ציוני הנחמדות:  
6  
6  

בדקו האם הציון הממוצע של נועה ב3 התכונות C, E וA  הוא גדול לפחות כמו הציון של הממוצע במדגם הכולל עבור שלוש תכונות אלה, על סך 4.23. 


```{r comp_op_ex, exercise=TRUE, exercise.eval = FALSE}

```

```{r comp_op_ex-check}
grade_this({if (identical(.result, TRUE)) {
    pass("מעולה! הציון שלה גדול מהממוצע")
}
  fail("עוד לא הגעתם לתשובה!")
})
```

```{r comp_op_ex-solution}
(((3 + 8-6)/2)+((6 + 8-2)/2) + (6+6)/2)/3 >= 4.23

```

האם היחס שבין הציון הממוצע בתכונה C לתכונה E (חלוקה) קטן מ0.5?


```{r comp_op_ex2, exercise=TRUE, exercise.eval = FALSE}

```

```{r comp_op_ex2-check}
grade_this({if (identical(.result, TRUE)) {
    pass("נכון! היחס קטן מחצי")
}
  fail("עוד לא הגעתם לתשובה!")
})
```

```{r comp_op_ex2-solution}
((3 + 8-6)/2)/((6 + 8-2)/2) < 0.5

```

חשבו את הציון הממוצע של נועה ב־E, העלו אותו בריבוע, ובדקו האם הוא גדול מ־25.

```{r comp_op_ex3, exercise=TRUE, exercise.eval = FALSE}

```

```{r comp_op_ex3-check}
grade_this({
    if (identical(.result, FALSE)) {
        pass("נכון! הציון בריבוע אינו גדול מ־25.")
    } else {
        fail("משהו לא הסתדר – בדקו שוב את חישוב הממוצע, החזקה, וההשוואה.")
    }
})
```

```{r comp_op_ex3-solution}
((3 + (8 - 6)) / 2)^2 > 25

```

## פעולות לוגיות
כדי לשלב שתי פעולות יחסיות נוכל להשתמש בפעולות לוגיות (Logical operator). אשר יקחו שני היגדים יחסיים ונוכל לברר האם שניהם מתקיימים, אחד מהם או אף אחד מהם.  
`&` - גם  
`|` - או  
`!` - לא  

פעולות לוגיות מאפשרות לנו ליצור תנאים מורכבים יותר. למשל, נוכל לבחון האם נועה היא מצפונית (6 בממוצע) לפחות כמו שהיא נחמדה (6 בממוצע) וגם לפחות כמו שהיא מוחצנת (2.5 בממוצע)?

```{r}
6 >= 6 & 6 >= 2.5
```
הקוד הדפיס לנו את התוצאה TRUE, כלומר, התשובה היא כן.   


`איך נדע לבחור בתנאי הנכון?` 
כאשר אנחנו נתקלים בבעיה לוגית תמיד כדאי לשאול את עצמינו האם ההשוואה הרצויה היא תנאי 1 `וגם` תנאי 2 ? אם כן - נשתמש ב&.  
אם ההשוואה הרצויה היא תנאי 1 `או` תנאי 2 נשתמש ב`|`.  

כדי להדגים את השימוש באופרטור | ננסח את ההשוואה הבאה:
האם ציון המוחצנות של נועה גדול מציון הנעימות או המצפוניות שלה?  

```{r}
2.5 > 6 | 2.5 > 6
```

הפעם קיבלנו התוצאה של FALSE. כלומר, התשובה היא כמובן לא, הדוגמה הפשוטה הזאת יכולה לעזור לנו להבין את העיקרון.    

הפעולה הלוגית האחרונה היא פעולת שלילה (not). הוספה של ! לפני השוואה יחסית או לוגית תהפוך את תוצאת ההשוואה. כך שאם התוצאה המקורית הייתה FALSE היא תהפוך לTRUE ולהפך.  
בואו נראה מה קורה להשוואה הקודמת שביצענו עם ההוספה של האופרטור !.

```{r}
!(2.5 > 6 | 2.5 > 6)
```

תשובה של כן הופכת ללא (ולהפך). זה יהיה שימושי בשיעורים הבאים :)

`תרגיל`  

יש לנו הזדמנות לוודא אם התשובות של נועה לתכונת המצפוניות (C) חריגות (כלומר אם יש פער גדול בין שתי התשובות לשאלות המודדות את אותה תכונה). נוכל לעשות זאת על ידי השוואה של התשובות שלה לציון החציוני. למשל אם הציון שלה בשאלה הרגילה גבוה מהציון החציוני נצפה שהתשובה ההפוכה תהיה נמוכה ממנו. כתבו קוד שמבצע את ההשוואה הבאה:
האם התשובה הראשונה של נועה (לשאלה רגילה) גדולה מ4 וגם התשובה ההפוכה גדולה או שווה ל4?   
תזכורת, התשובות לשאלה:   
שאלה רגילה - 6  
שאלה הפוכה - 2  

```{r logi_op_ex1, exercise=TRUE, exercise.eval = FALSE}

```

```{r logi_op_ex1-solution}
6 > 4 & 2 >= 4
```

```{r logi_op_ex1-check}
grade_this({if (identical(.result, FALSE)) {
    pass(" נכון! התוצאות לא חריגות שכן השאלה הרגילה היא מעל הציון החציוני והשאלה ההפוכה מתחת")
}
  fail("עד לא.. האם השתמשתם בפעולות ההשוואה ובפעולה הלוגית הנכונה?")
})

```



נמשיך לוודא את תשובותיה של נועה. לרוב אנשים שמקבלים ציונים מעל החציון (4) בתכונות "מצפוניות" (C) ו"נחמדות" (A) מקבלים גם ציון מתחת לחציון בתכונה "מוחצנות" (E). ולהיפך, אם ציון המוחצנות (E) מעל החציון ציוני המצפוניות (C) והנחמדות (A) לרוב מתחת לחציון.
בדקו האם  אחד התנאים מתקיים בנתוניה של נועה.

תוכלו להשתמש בציונים הממוצעים לפתרון שאלה זאת:  
E = 2.5

A = 6

C = 6

```{r logi_op_ex2, exercise=TRUE, exercise.eval = FALSE}

```

```{r logi_op_ex2-solution}
(6 > 4 & 6 > 4 & 2.5 < 4) | (6 < 4 & 6 < 4 & 2.5 > 4)
```

```{r logi_op_ex2-check}
grade_this({if (identical(.result, TRUE)) {
    pass(" מעולה אחד התנאים מתקיים, כלומר הנתונים של נועה עולים בקנה אחד עם המוכר בספרות")
}
  fail("זאת לא התשובה, שימו לב שצריך לבצע פה השוואה של וגם וגם של או?")
})

```


חוקרת במחלקה עורכת מחקר על אישיות מרצה.
החוקרת רוצה לכלול בניסוי רק משתתפים שמתקיימים אצלם שני תנאים הבאים ביחד:

ציון המוחצנות (E) שלהם נמוך מהממוצע בכיתה (4.3).

סכום הציונים הממוצעים שלהם ב"נעימות" (A) וב"מצפוניות" (C) גדול מהציון הממוצע במוחצנות בריבוע.

הנתונים של נועה:

E = 2.5

A = 6

C = 6

האם ניתן להכליל את נועה במחקר? בדקו זאת באמצעות פעולות לוגית והשוואתיות.

```{r logi_op_ex3, exercise=TRUE, exercise.eval = FALSE}

```

```{r logi_op_ex3-solution}
2.5 < 4.3 & (6+6) > 2.5^2
```

```{r logi_op_ex3-check}
grade_this({if (identical(.result, TRUE)) {
    pass("נכון! נועה שלנו יכולה להשתתף בניסוי")
}
  fail("עד לא.. האם השתמשתם בפעולות ההשוואה ובפעולה הלוגית הנכונה?")
})

```


## משתנים
עד כה השתמשנו בקוד כדי להדפיס תוצאות של תרגילים. פעולת ה`הדפסה` היא פעולה חשובה ושימושית, אבל מדובר בפעולה שרק מציגה את התוצאות, אך לא שומרת אותן בזכרון של המחשב.  
כדי לבצע פעולות מורכבות יותר באמצעות התוכנה נצטרך להשתמש במשתנים. המשתנים מאפשרים לנו לאחסן מידע בזכרון של המחשב. כך אנחנו יכולים להשתמש בערכים מסוימים שוב ושוב וגם לעדכן אותם בהתאם לצרכינו.  
כאשר מבצעים `השמה של משתנה` אנחנו צריכים לפרט שני מרכיבים:  

1. השם של המשתנה  
2. הערך של המשתנה   


### פעולות השמה
סוג הפעולה האחרון שנלמד להכיר הוא פעולת ההשמה.  פעולת השמה היא קביעה שמשתנה מסוים שווה לערך כלשהו. בR יש שתי דרכים לבצע פעולת השמה:

`->`   
או   
`=`  
אפשר להשתמש בשתי הדרכים ואין ממש הבדל בינהן.  

לדוגמה בקוד מטה מופיעה דוגמה פשוטה של השמת משתנה. שם המשתנה בדוגמה זאת הוא a והערך שלו 5.  
כלומר, החץ מצביע אל שם המשתנה.  


שנו את שם המשתנה לשם שלכם ואת הערך לגיל שלכם והריצו את הקוד והגישו את התשובה.

```{r var_example_a, exercise=TRUE, exercise.eval = FALSE}
a <- 5
```

```{r var_example_a-check}
grade_this({if (.result>=18) {pass("מצויין! בואו נמשיך הלאה ונלמד על משתנים.")
}
  fail("האם שיניתם את ערך המשתנה לגיל שלכם? אנחנו מצפים לתלמידים מעל גיל 18...")
})

```

ניתן גם להשתמש במשנים וגם בערכים כדי להשתמש בכל הפעולות שהכרנו בחלק הקודם. גם אתם הרגשתם שקשה לעקוב אחרי כל הציונים השונים של נועה? בואו נעשה סדר באמצעות משתנים:


קודם כל נגדיר שני משתנים שהם התשובות של נועה לשתי השאלות המוחצנות (E) .

```{r}
E1 <- 3
E2 <- 6
```

שימו לב, הגדרת המשתנים בפני עצמה לא מציגה כלום  (כלומר אין הדפסה לקוד זה) אבל מעכשיו המשתנים E1 וE2 קיימים בזכרון של המחשב.  

כעת נוכל לבצע פעולות פשוטות עם המשתנים שהגדרנו, למשל לחשב את הסכום

```{r}
E1 + E2
```
שימו לב, ביצוע החישוב באמצעות שני משתנים כן גרם להדפסה, זאת כיוון שהפעם לא שמרנו את התוצאה במשתנה כלל.  

אנחנו יכולים כמובן לאחסן את התוצאה של חישוב במשתנה. למשל, את הציון המתוקן של השאלה השניה (ההפוכה) במשתנה חדש

```{r}
corrected_E2 <- 8-E2
# השורה הראשונה לא תציג את התוצאה
# כדי להציג אותה נדפיס
# כתיבה של שם המשתנה תדפיס את הערך שהוא מכיל

corrected_E2

# נחשב את הציון הכללי
E <- (E1 + corrected_E2)/2

# ושוב נדפיס אותו
E

```

אפשר גם לבצע השוואות בין משתנים, למשל האם הציון המתוקן קטן מהציון המקורי ולאחסן את תוצאות ההשוואה במשתנה חדש.

```{r}
is_bigger <- E2 < corrected_E2
is_bigger
```

האם שמתם לב שהמשתנה is_bigger והמשתנה E שונים משני המשתנים הראשונים שהגדרנו (E1 וE2)?

```{r var_question, echo = FALSE}
question(" מה לדעתכם ההבדל הקריטי ביניהם?",
         answer("סוג הערך שמאוחסן במשתנים ", correct = TRUE),
         answer("השם של המשתנים הללו כולל מילה שלמה ולא כולל מספר", message = "האמנם זהו הבדל בין המשתנים, אך זה לא הבדל קריטי, שם המשתנה יכול להיות כל דבר. מילה מלאה, אות ואפילו כמה מילים או מספרים. כל עוד לא כללנו רווח."),
         answer("העובדה שמשתנים אלה הם תוצאה של חישוב ולא ערך", message = "זה אכן נראה כמו הבדל חשוב, אבל האמת היא שלשפת התכנות לא אכפת איך נוצר המשתנה, היא רק מאחסנת את הערך הסופי שמתקבל"),
         answer("שם המשתנים כולל קו תחתון", message = "בהמשך להסבר של תשובה 2, קו תחתון היא אחד הדרכים ליצור משתנה שמורכב מכמה מילים ללא רווחים"),
         allow_retry = TRUE
)
```

בחלק הבא נלמד על סוגי המשתנים השונים 

### משתנים מספריים  
עד כה הגדרנו רק משתנים מספריים. בשפת התכנות R משתנים מספריים נקראים באנגלית `Numeric variable`
בשפת התכנות יש שני סוגים של משתנים מספריים:  

`Integer` -  הערך של המשתנה הוא מספר שלם  
`Float` - הערך של המשתנה הוא שבר  


```{r numeric_question, echo = FALSE}
question(" אם ניצור משתנה שהוא חיבור בין מספר שלם לשבר מה יהיה סוד המשתנה?",
         answer("Float ", correct = TRUE, message = "נכון, חיבור בין שבר למספר שלם יתן לנו תמיד שבר."),
         answer("Integer", message = "לא נכון"),
         allow_retry = TRUE
)
```


### משתנים מסוג טקסט
לא כל המשתנים בשפת התכנות חייבים להיות מספריים. ניתן גם להשתמש בטקסט בשפת התכנות. גם כאן יש שתי אפשרויות:  

`Character` - משתנה מסוג טקסט באופן כללי בR. המונח מתאר אותיות בודדות באופן כללי.  
`String` - משתנים המכילים מספר אותיות נקראים string כיוון שהם מהווים "שרשרת" של אותיות. 

`שימו לב`  
שפת התכנות R היא שפה שנקראית Case sensitive כלומר היא מבחינה בין אותיות קטנות (למשל a) לגדולות (למשל A).  

`גם שמות המשתנים מורכבים מטקסט, איך נבדיל בניהם?`  
שפת התכנות מבדילה בין שמות המשתנים לערך שהינו טקסט על ידי שימוש בגרשיים כפולים `""` או בודדים `''`.  
לדוגמה כאן יצרנו משתנה המכיל את השם של המשתתפת שלנו  

```{r}
name <- "Noa"

# נדפיס אותו
name

```

`תרגיל`  
צרו משתנה הנקרא full_name המכיל את השם המלא של נועה, נועה כהן (Cohen). הקפידו לכלול רווח בין השם הפרטי לשם המשפחה.
```{r string_ex1, exercise=TRUE, exercise.eval = FALSE}

```

```{r string_ex1-check}
grade_this({if (.result=="Noa Cohen") {pass("מצויין!")
}
  fail("האם שמתם לב לאותיות גדולות? שימו לב כי מצופה לרווח בין השם הפרטי לשם המשפחה")
})
```


### משתנים לוגיים
או באנגלית `Logical varaible` הוא משתנה עם ערך לוגי, כלומר משתנה שיכול לקבל שני ערכים, נכון (TRUE) או לא נכון (FALSE).  

זוכרים שלמדנו ששפת התכנות בR רגישה לסוג האותיות (קטנות מול גדולות)? אז בהתאם היא מצפה שהמילים במשתנים הלוגים יהיו כתובים באותיות גדולות ללא גרשיים.  

יש למשתנים אלה שימושים רבים, שאת חלקם נכיר בהמשך. כרגע נתמקד בכך שאם נאחסן בתוך משתנה השוואה בין שני משתנים המשתנה שיווצר יהיה משתנה לוגי שיכיל את תוצאת ההשוואה.

 כבר ראינו דוגמאות להשוואות בין מספרים אבל גם משתני טקסט ניתן להשוות זה לזה. האמנם אין משמעות להאם טקסט 1 גדול או קטן מטקסט 2, אבל האם מחרוזת אחת זהה (או דווקא לא זהה) לאחרת יכולה להיות פעולה מאוד שימושית. 

`תרגיל`  
חוקרת קיבלה נתונים של שתי משתתפות, היא בודדה את השם של כל משתתפת במשתנה נפרד, הנתונים לכם מטה. בדקו האם מדובר באותה משתתפת?

```{r str_comp_ex, exercise=TRUE, exercise.eval = FALSE}
name_1 <- "Noa Cohen"
name_2 <- "Noa Levi"
```

```{r str_comp_ex-check}
grade_this_code(correct =  "נכון מאוד. שמתם לב,
שבמתשנה מסוג זה ניתן לכלול רווחים, וכמעט תו אחר? 
ההשוואה כמובן מביאה לתוצאה שלילית שכן המחרוזות לא זהות"  , incorrect = "נסו שוב, לא הוספתם את ההשוואה הנכונה"  
)
```

```{r str_comp_ex-solution, exercise.reveal_solution = FALSE}
name_1 <- "Noa Cohen"
name_2 <- "Noa Levi"

name_1 == name_2
```

עכשיו נעלה קצת ברמת הקושי. 
נתונות לכם שתי המחרוזות הבאות:  
name1 = "Noa Cohen"  
name2 = "Noa cohen"  

```{r str_question, echo = FALSE}
question(" מה היינו מקבלים אם היינו משווים בין המחרוזות?",
         answer("נקבל תשובה חיובית כמובן!", message = "האמנם זה אותו השם, אבל שפת התכנות שלנו רגישה להבדל בין אותיות קטנות לגדולות באנגלית."),
         answer("ברור שנקבל לא!",correct = TRUE, message ="נכון! שפת התכנות רגישה לאותיות קטנות מול גדולות ולכן למרות שזה אותו שם נקבל תשובה שלילית"),
         allow_retry = TRUE
)
```
  

`תרגיל`  
נתונים לכם כאן נתונים נוספים לגבי שתי המשתתפות בשאלון. בדקו האם המקצוע הלימודים שלהן שונה זו מזו וגם הציון הממוצע שלהן בתכונה A שונה זו מזו ושמרו את התוצאה במשתנה.

```{r log_ex, exercise=TRUE, exercise.eval = FALSE}
Major_1 <- "Psychology"
Major_2 <- "Computer Science"
mean_A_1 <- 6
mean_A_2 <- 1


```

```{r log_ex-check}
grade_this({if (.result==TRUE) {pass("כל הכבוד! יצרתם שתי השוואת בין סוגי משתנים")
}
  fail("עוד לא הגעתם לתוצאה הנכונה")
})

```

```{r log_ex-solution, exercise.reveal_solution = FALSE}
Major_1 <- "Psychology"
Major_2 <- "Computer Science"
mean_A_1 <- 6
mean_A_2 <- 1

# חשבו כאן את log_1
logical <- Major_1 != Major_2 & mean_A_1 != mean_A_2

```

## תרגיל מסכם

נעבור לתרגול יחד את כל מה שלמדנו בלומדה זאת.  

בתרגיל זה נפגוש את ציוניה של נועה בשני מדדים נוספים, המבטאים את שתי התכונות האישיותיות הנותרות.

מדדים אלו כוללים את תכונת ה"פתיחות לחוויות" המסומנת באות O ואת תכונת ה"נוירוטיות" המסומנת באות N. כל אחת מאלו נמדדה על ידי 2 שאלות, אחת עם דירוג רגיל ואחת בעלת דירוג הפוך.



`שאלה 1`  

במדד O נועה דיווחה `5` גם בשאלה הרגילה וגם בשאלה ההפוכה.  
צרו שני משתנים (`O1` ו ``O2`) המכילים את ערכי התשובות של נועה עבור שתי השאלות במדד O והיעזרו במשתנים אלו כדי לחשב את ממוצע הדירוגים של נועה עבור מדד O (שימו לב להפוך את ציוני השאלה ההפוכה לפני החישוב).

```{r final_ex1, exercise=TRUE, exercise.eval = FALSE}

```

<div id="final_ex1-hint">
**Hint:** נתקעתם? בואו נפרק את התרגיל.  
קודם כל צרו שני משתנים המכילים את התשובות של נועה. שימו לב כי שם המשתנה צריך להיות ללא גרשיים (גרשיים מעידים על ערך מסוג טקסט ולא על שם של משתנה).
לאחר שיש לנו את המשתנים הצעד הבא יהיה להפוך את השאלה ההפוכה (שאלה מספר 2), תוכלו ליצור משתנה חדש המכיל את הערך המתוקן.   
לא זוכרים איך להפוך? חזרו לפעולות אריתמתיות.  
לאחר שיש לכם את זה כל מה שנשאר זה ליצור משתנה חדש שיקרא O ומכיל את הממוצע של הציון הרגיל ואת הציון המתוקן בשאלה ההפוכה. תוכלו למצוא דוגמה לביצוע חישוב זה בדיוק תחת פעולות אריתמתיות.
</div>


```{r final_ex1-solution, exercise.reveal_solution = FALSE}

O1 <- 5
O2 <- 5

ave_O <- (O1 + 8-O2)/2
```

```{r final_ex1-check}
grade_this({if (identical(.result, (5 + 8-5)/2)) {
  pass("מעולה! הציון הממוצע של נועה בתכונת הפתיחות הוא 4.")
}
  fail("זאת לא התשובה הנכונה.. האם הפכתם את הפריט ההפוך? יכול להיות ששכתם לשים סוגריים? בדקו את הצבת הנתונים במשתנים !")
})
```

`שאלה 2`  

הציון הממוצע של נועה במשתנה הנוירוטיות (N) הוא 3.5.
משתנה זה קשור באופן החזק ביותר להסתברות לפתח הפרעות פסיכולוגיות שונות. 
עורכי המחקר רצו לפתח מדד שישקף את רמתו הדומיננטיות שלו מול רכיב N ביחס לרכיבי האישיות האחרים.  
החוקרים קראו למדד "N Ratio".
חשבו את המדד על ידי חלוקת הציון הממוצע בN בסכום של הציון הממוצע בכל שאר התכונות.

השלימו את החישוב בקוד מטה:
```{r final_ex2, exercise=TRUE, exercise.eval = FALSE}
C <- 6
A <- 6
E <- 2.5
# כאן יש להוסיף את הנתונים על פתיחות ונוירוטיות

# כאן חשבו את הסכום של כל מדדי האישיות שאינן נוירוטיות
big_5_sum <-
  
# ולבסוף חשבו את היחסים
N_ratio <- 
  
```

<div id="final_ex2-hint">
**Hint:** צריכים עזרה?  
קודם כל הוסיפו שני משתנים לנתונים - את O שחישבתם בתרגיל הקודם ואת N שנתון לכם בתחילת השאלה. 

לאחר מכן אכסנו במשתנה שנתון לכם ונקרא big_5_sum את הסכום של כל התכונות מלבד התכונה N. זאת כיוון שבשאלה יחס הN מחושב על ידי חלוקה של הציון בN בכל בסכום של כל שאר הציונים.  

לבסוף חשבו את יחס הN. כל מה שדרוש הוא חלוקה של שני משתנים שכבר קיימים בשלב זה. לא זוכרים איך לבצע חילוק? חזרו לפעולות אריתמתיות.
</div>

```{r final_ex2-solution, exercise.reveal_solution = FALSE}
C <- 6
A <- 6
E <- 2.5
# כאן יש להוסיף את הנתונים על פתיחות ונוירוטיות
O <- 4
N <- 3.5
# כאן חשבו את הסכום של כל מדדי האישיות יחד
big_5_sum <- O + C + E + A

# ולבסוף חשבו את החלק של נוירוטיות ברכיבי האישיות
N_ratio <- N/big_5_sum
```

```{r final_ex2-check}
grade_this({if (identical(.result, 3.5/(6+6+2.5+4))) {
  pass("כל הכבוד! חישבתם את הN_Ratio בצורה מדוייקת!")
}
  fail("זאת לא התשובה הנכונה.. יכול להיות ששכתם לשים סוגריים? קראו שוב את ההנחיות ןבדקו את הצבת הנתונים במשתנים. בדקו את כיוון החלוקה")
})
```

`שאלה 3`  
חוקר במחלקה מנהל מחקר על דיכאון לאחר לימודים אקדמיים. הוא מעוניין לגייס למחקר סטודנטים הלומדים פסיכולוגיה או עבודה סוציאלית לקבוצת מחקר של סטודנטים בעלי נוירוטיות גבוהה (עם יחס נוירוטיות של לפחות 0.2). כתבו קוד שבודק ומדפיס האם נועה שלנו מתאימה לקבוצת המחקר

```{r final_ex3, exercise=TRUE, exercise.eval = FALSE}
major <- "Psychology"
# הציבו את התשובה שלכם מהתרגיל הקודם
N_ratio <-
# כאן בצעו את ההשואה והדפיסו את התוצאה
  
```

<div id="final_ex3-hint">
**Hint:** יש לכם שאלות?  
קודם כל הוסיפו את התוצאה שלכם מהתרגיל הקודם למשתנה המתאים.
עכשיו נבצע את ההשוואה. בדקו בהשוואה האם נועה לומדת פסיכולוגיה או(!) עבודה סוציאלית וגם (!) בעלת יחס N גבוה מ0.2?

לא זוכרים איך לתרגם את ההשוואות לקוד? חזרו לחלק על פעולות לוגיות!

שימו לב לסוגריים! ההשוואה של תחום הלימוד צריכה להתבצע לפני ההשוואה בין התוצאה של השוואת תחום הלימוד 
</div>

```{r final_ex3-solution, exercise.reveal_solution = FALSE}
major <- "Psychology"
# הציבו את התשובה שלכם מהתרגיל הקודם
N_ratio <- 0.18
# כאן בצעו את ההשואה והדפיסו את התוצאה
is_high_N <- N_ratio <= 0.2 & (major == "Psychology" | major == "Social work")
is_low_N
```

```{r final_ex3-check}
grade_this({if (identical(.result, FALSE)) {
  pass("קיבלנו תשובה חיובית, לנועה יש ציון של פחות מ0.2 ולכן היא לא מתאימה לקבוצת המחקר")
}
  fail("קראו שוב את ההנחיות, עוד לא הגעתם לתשובה הנכונה")
})
```

## משוב

בחלק זה נבקש את המשוב שלכם על הלומדה. אנא ענו בכנות וביסודיות, על מנת שנוכל להשתפר. התשובות ישמרו בצורה אנונימית ולא ישפיעו על בדיקת התרגיל עצמה.  
שימו לב: ברוב השאלות אין אפשרות לשנות את התגובה לאחר לחיצה על כפתור ההגשה.

```{r survey_q1, echo = FALSE}
question("מה הייתה רמת הקושי של הלומדה עבורך? יש להתייחס לכמה הבנת את החומר ולרמת הקושי של התרגול עבורך. ",
         type="single",
         answer("קלה"        ,correct=TRUE,  message = "תגובתך: קלה"),
         answer("בינונית"    ,correct=TRUE,  message = "תגובתך: בינונית"),
         answer("מאתגרת"     ,correct=TRUE,  message = "תגובתך: מאתגרת"),
         answer("קשה"        ,correct=TRUE,  message = "תגובתך: קשה"),
         answer("קשה מאוד"   ,correct=TRUE,  message = "תגובתך: קשה מאוד"),
         correct = "",
         incorrect = "",
         allow_retry = TRUE
)
```

```{r survey_q2, echo = FALSE}
question("עד כמה החומר הנלמד ביחידה הועבר בצורה ברורה לדעתך?",
         type="single",
         answer("במידה רבה"      , correct=TRUE,  message = "תגובתך: במידה רבה"),
         answer("במידה סבירה" ,    correct=TRUE,  message = "תגובתך: במידה סבירה"),
         answer("באופן חלקי"  ,    correct=TRUE,  message = "תגובתך: באופן חלקי"),
         answer("במידה מעטה"  ,correct=TRUE,  message = "תגובתך: במידה מעטה"),
         answer("במידה מעטה מאוד",correct=TRUE,  message = "תגובתך: במידה מעטה מאוד"),
         correct = "",
         incorrect = "",
         allow_retry = TRUE)
```

```{r survey_q3, echo = FALSE}
question("כמה זמן לקח לך לפתור את הלומדה בערך?",
         type="single",
         answer("פחות משעה"      , correct=TRUE,  message = "תגובתך: פחות משעה"),
         answer("שעה" ,    correct=TRUE,  message = "תגובתך: שעה"),
         answer("שעתיים"  ,    correct=TRUE,  message = "תגובתך: שעתיים "),
         answer("שלוש שעות"  ,correct=TRUE,  message = "תגובתך:  שלוש שעות"),
         answer("ארבע שעות ומעלה",correct=TRUE,  message = "תגובתך:  ארבע שעות ומעלה "),
         correct = "",
         incorrect = "",
         allow_retry = TRUE)
```

```{r survey_q4, echo = FALSE}
question(" אילו נושאים בלומדה היו קשים במיוחד עבורך?",type="learnr_checkbox",
         answer("פעולות אריתמטיות", correct=TRUE),
         answer("פעולות השוואה", correct=TRUE),
         answer("פעולות לוגיות", correct=TRUE),
         answer("משתנים", correct=TRUE),
         answer("תרגיל מסכם", correct=TRUE),
         correct = "תשובתך התקבלה. תודה על הפירוט",
         try_again = "תשובתך התקבלה. תודה על הפירוט",
         incorrect = "תשובתך התקבלה. תודה על הפירוט",
         allow_retry = TRUE
)
```

```{r survey_q5, echo = FALSE}
  question_text(
    " מצאת טעות? ספר/י לנו עליה (יש להקפיד ולציין את שם הפרק)",
    incorrect = "תשובתך התקבלה!",
    correct = "תשובתך התקבלה!",
    answer(text = "asd",message="DFGDF"),
    answer(text = "",correct = TRUE),
    rows = 10,
    trim = FALSE
  )
```

```{r survey_q6, echo = FALSE}
  question_text(
    "יש לך עוד משהו לספר לנו? נשמח לשמוע משוב מפורט, הצעות כלליות או ספציפיות לפרק מסוים לגבי יחידה זאת.",
    incorrect = "תשובתך התקבלה",
    correct = "תשובתך התקבלה",
    answer(text = "asd",message="DFGDF"),
    answer(text = "",correct = TRUE),
    rows = 10,
    trim = FALSE
  )
```


## הגשת התרגיל
  סיימת? מעולה! עכשיו הגיע הזמן להגיש את התרגיל.  
  יש ללחוץ על הכפתור:  Generate  
  להעתיק את הטקסט שמופיע בחלון למטה ולהגישו במודל  
  בהצלחה!

```{r context="server"}
learnrhash::encoder_logic()
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui()
```


הדביקו את קוד הhash בתוך הגרשיים בפונקצייה הבאה כדי לוודא שעניתם על כל השאלות והתרגילים בלומדה.

```{r hash_check,  exercise = TRUE, exercise.eval = FALSE}

check_hash_code("")

```



