---
title: "Lesson 07: Coding resources (Prerequisite for the summary Exercise)" 
author: "Copyright 2025 Psychology Department. Hebrew University of Jerusalem. All rights reserved"

output:
  learnr::tutorial:
    progressive: true,
    allow_skip: true,
    
description: >
 <div style='direction: rtl;'>
 השיעור נוגע בעזרים חיצוניים שניתן ומומלץ להשתמש בהם בעת ניתוח נתונים בR. נדבר על תיעוד נלווה, בלוגים ופורומים אינטרנטיים ושימוש בבינה מלאכותית.
 </div>
     
runtime: shiny_prerendered  
editor_options: 
  markdown: 
    wrap: 72
---


<style>
h1, h2, h3, h4, h5, h6 {
  direction: rtl;
}
p {
  direction: rtl;
}

.text-block1 {
  direction: rtl;       /* Set text direction to right-to-left */
  text-align: right;
  background-color: #f0f0f0; /* Light grey background */
  padding: 10px;
  border-radius: 5px;
  border: 1px solid #ddd; /* Light border */
  margin: 10px 0; /* Space around the block */
}
</style>


```{r setup, include=FALSE}
library(learnr)
library(dplyr)
library(gradethis)

knitr::opts_chunk$set(error = TRUE, warning = TRUE) 

{
exercises_df = data.frame(
  exercises = c("rdocumentation",
                "separate",
                "warnings"),
  
  hebrew = c("תיעוד נלווה",
             "שימוש בAI: בירור לגבי פונקציות",
             "שימוש בAI: פתרון תקלות")
)


check_hash_code   = function(hash){
    exercises = exercises_df$exercises
    response_table = learnrhash::decode_obj(hash)
    if (length(response_table)==0){return("Invalid hash code")}
    correct = response_table %>% filter(label %in% exercises, correct) %>% pull(label)
    incorrect = setdiff(exercises, correct)
    if (length(incorrect) == 0){return("Great work! Hash code is valid")}
    else{
      hebrew = exercises_df  %>% filter(exercises %in% incorrect) %>%pull(hebrew)
      print("The following exercises are missing or incorrect:")
      print(hebrew)
    }

} 
}# hash checker

```

```{r prepare-df1}
set.seed(111)
df1 = data.frame(subject = rep(1:5,each = 4),
                 trial   = rep(1:4,5),
                 score   = round(rep(1:4,5) + rnorm(20, 20,5)) %>% as.character())

```

```{r prepare-df2}

df = data.frame(session = c("s01_t1","s01_t2","s01_t3","s02_t1","s02_t2"),
                score = c(45,65,57,71,52))

```

## מבוא

ביחידה זו נלמד על מקורות מידע שימושיים לכתיבת קוד בR ופתרון תקלות.


אנחנו נתייחס לשלוש סוגי מקורות מומלצים:

**1)** תיעוד נלווה של פונקציות\
**2)** בלוגים ופורומים אינטרנטיים\
**3)** בינה מלאכותי

## תיעוד נלווה

המשאב הראשון והוותיק ביותר שעומד לרשותנו הינו התיעוד הנלווה של פונקציות
(Documentation). לכל חבילה בR מוצמד תיעוד מתאים - המפרט עבור כל פונקציה
את מטרתה, את הנתונים השונים שהיא יכולה או צריכה לקבל כדי לפעול באופן
תקין ומספר דוגמאות שימוש.

ניתן לגשת לתיעוד זה בשני דרכים:

### באמצעות הפקד `?`

כשנריץ סימן שאלה מלווה בשמה של פונקציה - יפתח עבורנו חלון העזרה "help"
המובנה בתוך Rstudio ובתוכה יופיע התיעוד הרלוונטי לפונקציה.

למשל - במידה ונרצה לברר איך משתמשים בפונקציה `round` נריץ את השורה הבאה:

```{r , eval = FALSE}
?round
```

הריצו את שורת הקוד בRstudio ועיינו בתיעוד הנלווה.

תוכלו לראות שהתיעוד מציג את אופן השימוש בפונקציה: היא מקבלת כאיבר ראשון
את X - שהוא האובייקט עליו נרצה להפעיל את הפונקציה, וכאיבר שני את מספר
הספרות אחרי הנקוב העשרונית אליהם נרצה לעגל (digits). במקרה זה התיעוד
כולל גם פירוט על מספר פונקציות דומות.

### Rdocumentation

התיעוד שמוצג בחלון העזרה מופיע גם באתר האינטרנט
[Rdocumentation](https://www.rdocumentation.org/ "קישור לאתר").

המידע באתר אינו שונה מחלון העזרה, אבל ההצגה שלו על פני דף שלם בדרך כלל
נוחה יותר מאשר התצוגה בחלונית.

בשני המקרים - המידע שמסופק הוא מאוד טכני ולא תמיד קל להבנה. הוא שימושי
בעיקר כדי להיזכר בנתונים שפונקציה אמורה לקבל וכדי לראות דוגמאות שימוש
בסיסיות.

`תרגיל:` חפשו ברשת או בחלונית העזרה את התיעוד הנלווה של הפונקציה
`substr`, המשמשת לשליפה של חלק מתוך מילה / משתנה תווי (character).

היעזרו בפונקציה זו על מנת לשלוף את שלושת התווים הראשונים בקוד המזהה של
הנבדקים ושמרו אותם בעמודה חדשה בשם "short_ID":

```{r rdocumentation, exercise=TRUE, exercise.eval = FALSE, exercise.setup = "prepare-df1"}

df = data.frame(ID = c("001p","002f","003f","004f","005p","006f"),
                score = c(45,65,57,71,52,43))

df = df %>% mutate(short_ID = substr(___________))

print(df$short_ID)

```

```{r rdocumentation-check}

df = data.frame(ID = c("001p","002f","003f","004f","005p","006f"),
                score = c(45,65,57,71,52,43))
  
  
grade_result(
  pass_if(~identical(.result,substr(df$ID,1,3) )),
         
  correct="יפה מאוד!", 
  incorrect = "נסו שוב"
)
```

## בלוגים ופורומים אינטרנטיים

משאב נוסף, שהינו מעט יותר ידידותי למשתמש, כולל מידע שהועלה על ידי
משתמשים לרשת.

בחלק מהמקרים מדובר בפורומים בהם אנשים העלו בעיות שהם נתקלו בהם וקיבלו
תשובות העשויות לעזור לנו בבעיות איתן אנחנו מתמודדים ובחלקם מדובר בבלוגים
מסודרים, בהם אנשים מעלים מדריכים מפורטים לגבי אופן השימוש בפונקציות
מוסימות ודרכים בהם אפשר לבצע ניתוחים מורכבים. הפירוט כאן לרוב יהיה עשיר
יותר, לפעמים בליווי דוגמאות או סרטונים רלוונטיים.

בנוסף, ישנם סיכומים מסודרים לפונקציות מחבילות שונות ואופן השימוש בהן שהועלו לרשת כקבצי PDF שנועדו לאפשר חיפוש מהיר ונוח של מידע.

את רוב הרשומות והשאלות בפורומים אפשר למצוא באמצעות חיפוש ברשת של הבעיה
איתה אתם מתמודדים, כך שאין חשיבות אמיתית להיכרות עם הבלוגים השונים מראש
אבל נמליץ כאן בכל זאת על בלוג אחד וכמה דפי מידע העשויים להיות שימושיים:

::: text-block1
[The R Graph Gallery](https://r-graph-gallery.com/ "קישור לאתר")- בלוג
המפרט על אופן היצירה של סוגים שונים של תרשימים, רובם באמצעות `ggplot`
ומספר חבילות המאפשרות לייצר תרשימים אינטראקטיביים.


**דפי מידע:**

   [פונקציות בסיסיות](https://iqss.github.io/dss-workshops/R/Rintro/base-r-cheat-sheet.pdf)
  
   [עריכת נתונים עם החבילות dplyr שילמדו ביחידה 5](https://nyu-cdsc.github.io/learningr/assets/data-transformation.pdf)
    
   [תרשימים בggplot](https://posit.co/wp-content/uploads/2022/10/data-visualization-1.pdf)
    
   
:::

## שימוש נכון בבינה מלאכותית

### מבוא

עד לאחרונה, השימוש ברשומות ופורומים היווה את המשאב העיקרי עבור למידה
ופתרון בעיות בR, אבל מאז העלייה של מנועי הבינה המלאכותית בשנים האחרונות
רבים עוברים לעשות שימוש במנועים כדוגמת chatGPT, Claude וכו'. היתרון
בשימוש בבינה מלאכותית הוא בכך שהם מאפשרים לשאול שאלות תכנותיות באופן
ישיר (חלקם גם תומכים בשאלות בעברית) ולקבל תשובה מפורטת המותאמת היטב למה
שאנחנו זקוקים לו. התוכנות מסוגלות גם לייצר קטעי קוד, קצרים או ארוכים,
כדי להדגים את התשובות שלהם ולתקן בעיות בקוד שלנו.

הכלים האלו אמנם מגבירים את היכולת שלנו לפתור בעיות ולבצע ניתוחים
מורכבים - אבל חשוב להכיר גם את המגבלות שלהם.

![](images/ai.jpg){width="100"}

בפרק הנוכחי נדבר על מספר שימושים מומלצים בבינה מלאכותית ונתעכב על כמה
אזהרות חשובות.

### בירור לגבי פונקציות

השימוש הפשוט ביותר בבינה מלאכותית, בו הוא גם בדרך כלל מצטיין, כולל
התייחסות אליו כאל מאגר מידע. במקום לנסות מילות מפתח שונות ולגלול ברחבי
הרשת עד שתמצאו את הפונקציה הספציפית שאתם צריכים, או רשומה נוחה לגבי אופן
השימוש בה - בינה מלאכותית בדרך כלל תדע לספק לכם תשובות מהירות ומבוססות לגבי
איזו פונקציה מבצעת את הפעולה שאתם מחפשים לבצע או לגבי הדרכים השונות להשתמש
בפונקציה נתונה.

::: text-block1
גם כאן יכולות להיות טעויות והתכנה לפעמים תמציא פונקציות או תשמשתש
בהן בצורה לא נכונה, אבל ברוב המקרים הפשוטים היא מספקת מענה מוצלח מאוד.
בכל אופן - נוכל לבדוק את התשובות שלה באופן מיידי, בין אם באמצעות הרצת
הפונקציה או באמצעות עיון בתיעוד הנלווה של הפונצקיה וברשומות קיימות לגביה.
:::

`תרגיל` : בררו איזו פונקציה יכולה לשמש אותנו כדי להפריד עמודה לשניים.

בטבלה הבאה ישנה עמודה בשם "session" המכילה מידע גם על מספר הנבדק וגם על מספר הסבב של אותו הנבדק. הפרידו אותה לעמודות
"subject"ולעמודה "trial" כך שכל עמודה תכיל את התוכן המתאים.

למשל - בשורה הראשונה נרצה שתחת העמודה "subject" יופיע "s01" ותחת העמודה
"trial" יופיע "t1".

```{r separate, exercise = TRUE, exercise.setup = "prepare-df2"}

print(df)

df = df %>% ________________
  
print(df)

```

```{r separate-check}

grade_result(
  pass_if(~identical(.result, 
                     
data.frame(session = c("s01_t1","s01_t2","s01_t3","s02_t1","s02_t2"),
                score = c(45,65,57,71,52))%>%
  separate(session, into = c("subject","trial"))
  
  )),
         
  correct="מעולה!", 
  incorrect = "נסו שוב"
)
```

### פתרון תקלות

אחד העיסוקים הכי מתסכלים בתכנות בכלל, ובניתוח נתונים בR בפרט, הוא שמדי
פעם אנחנו נתקלים בתקלות ובעיות בהרצה שלא תמיד ברור לנו איך עלינו לפתור
אותן. הפלט בR יכול להכווין אותנו לפונקציה שבה ארעה התקלה - אבל גם אז
לפעמים קשה להבין ממה התקלה נובעת.

מנועי בינה מלאכותית יכולים לעזור לנו לצאת מהבוץ בחלק ממקרים אלו. אם
נדביק את הפלט של השגיאה למנוע הבינה המלאכותית הוא יוכל להסביר לנו בפירוט
את המשמעות של השגיאה שחווינו. הוא לא תמיד יוכל להסיק את שורש הבעיה, אבל
הוא יוכל לתת לנו הבנה ראשונית של הבעיה לכל הפחות.

אם נצרף לפנייה שלנו גם את החלק הבעייתי בקוד שלנו, התוכנה תוכל להבין
בצורה טובה יותר את מקור הבעיה ובמקרים רבים גם להציע לנו איך לתקן אותה.

`תרגיל` הריצו את קטע הקוד הבאה והיעזרו במנוע בינה מלאכותית לבחירתכם כדי
להבין את מקור השגיאות ולתקן אותן. שימו לב - הקוד מכיל מספר שגיאות שונות.


```{r warnings, exercise = TRUE, exercise.setup = "prepare-df1"}
# df1:
# הטבלה מתארת את הציונים של נבדק בכמה חזרות שונות של אותה המטלה

# subject - המספר המזהה של הנבדק
# score  - ציון הנבדק
# trial   - מספר החזרה

df1 %>% head() %>% print() # הדפסת השורות הראשונות

df1_summary = df1 %>% 
  group_by(subject) %>%                      
  summarise(trial_amount = max(trial), # מספר החזרות שביצע
            mean_score = mean(score))       # הציון הממוצע 
               

mean_scores = df1_summery %>% pull(mean_score)

print(mean_scores)

```

```{r warnings-check}


grade_this({
  
  target  = df1 %>% 
    group_by(subject) %>%                      
    summarise(trial_amount = max(trial, na.rm = T), 
              mean_score = mean(as.numeric(score), na.rm = T))   %>%
    pull(mean_score)
  
  result = user_object_get("mean_scores")

   if (is.na(sum(result))){
     fail("האם הוקטור שהודפס נראה תקין?")
   } else if (identical(target, result)){
     pass("עבודה מצוינת!")
   }  else {
     fail("עדיין ישנה שגיאה בקוד")
   }
  
   }
  )
```

### יצירת קטעי קוד

שימוש נוסף, יעיל אך מעט בעייתי, בבינה מלאכותית הוא יצירה של קטעי קוד
שלמים. במקום להקדיש זמן יקר לכתיבת קוד, זיהוי שגיאות ותיקונים, אנחנו
לכאורה יכולים לבקש מהבינה המאכותית לכתוב את הקוד עבורנו.


יש בכך 3 בעיות עיקריות:

::: text-block1

**1) טעויות**

טעויות קטנות בקוד יגרמו עשויות לגרום לשגיאות - ובמידה והקוד לא נכתב על ידינו - יהיה
לנו מאוד קשה להבין את מקור הטעות ולתקן אותה. נוכל להיעזר בבינה המלאוכתית
בשביל תיקונים כאלו, אבל במקרים רבים נמצא את עצמנו עם שמיכת טלאים של קטעי
קוד שאנחנו לא מצליחים לתקן. זה קורה בעיקר בקטעי קוד שמבצעים יותר מצעד עיבוד אחד.


כפועל יוצא מכך - מומלץ להגביל את השימוש בקוד שנוצר על ידי בינה מלאכותית לפסקה או פעולה אחת בכל פעם. 
  
אפשר ורצוי לבקש מהתוכנה שתסביר מה כל
חלק עושה - אבל כל עוד לא הבנו מה כל שורה עושה, השימוש וההתבססות על קטעי
קוד כאלו הוא הימור בעייתי. זה אמנם חוסך זמן בטווח הקצר, אבל במקרים רבים יוסיף
לכם הרבה עבודה ותסכול בהמשך.

**2) קצרים בתקשורת**

במובנים רבים, האפשרות הזו יותר מסוכנת לנו מאשר יצירת קוד שלא עובד. במידה
והתכנה לא תבין את הפרטים של הניתוח שאתם מעוניינים לבצע - היא יכולה לספק לכם קוד
מפורט ותקין שמבצע ניתוח כלשהו ומספק לכם תוצאות, אבל אילו לא יהיו התוצאות של הניתוח שאותו רציתם לערוך. כלומר - בנוסף
להשגחה על תקינות הקוד מבחינה תכנותית, חשוב גם לוודא שהקוד באמת מבצע את
הניתוח שאנחנו מעוניינים בו.


**3) יצירת תלות**

בשלב זה אתם עוד לומדים להשתמש בתוכנה. כלומר - יש הרבה פונקציות ופרקטיקות
ניתוח בסיסיות שאתם עוד רוכשים. מצד אחד - אלו תחומים בהם הבינה המלאכותית
בדרך כלל שולטת היטב ותוכל לספק תשובות טובות, אבל מצד שני - אם לא תלמדו
אותם, תתרגלו ותתנסו בהם, ייקח הרבה זמן עד שתוכלו לבצע את ניתוחים אלו
בעצמכם. התלות בבינה מלאכותית בנושאים אלו תעכב ותסרבל את העבודה שלכם.

בשלב זה, וגם בהמשך העבודה עם בינה מלאכותית, מומלץ להשתמש בה רק כמורה -
שמסביר, מדגים ועוזר לכם למצוא את הטעויות שלכם ולתקן אותן - אבל לא מבצע
את הניתוחים במקומכם.

בקורס הזה אנחנו מצפים שתכתבו את הקוד כולו בעצמכם ושתבינו מה כל שורה מבצעת R.
:::


### קופיילוט ![](images/copilot.png){width="50" height="26"}

תצורה נוספת ושימושית במיוחד של בינה מלאכותית הינה Github Copilot. היישום
כולל תוסף לRstudio הזמין החל מהגרסאות האחרונות של התוכנה, ומציע השלמות
לקוד בזמן הכתיבה.

לתוסף יש גישה לקובץ שאתם עובדים עליו ולקבצים אחרים שכתבתם, ובאמצעותם הוא
מנסה לנבא את מטרת הקוד שאתם כותבים ולהשלים אותו בהתאם.

מאחר והתוסף לא דורש מאיתנו להעתיק את הקוד לממשק אחר ולנסח באופן מפורש את
השאלות שלנו הוא הרבה יותר מהיר ונח לשימוש.

נכון להיום, התוסף דורש מנוי בתשלום, אך הוא מוצע חינם לסטודנטים דרך תכנית
Github education.


## משוב

בחלק זה נבקש את המשוב שלכם על הלומדה. אנא ענו בכנות וביסודיות, על מנת שנוכל להשתפר. התשובות ישמרו בצורה אנונימית ולא ישפיעו על בדיקת התרגיל עצמה.

שימו לב: ברוב השאלות אין אפשרות לשנות את התגובה לאחר לחיצה על כפתור ההגשה.

```{r survey_q1, echo = FALSE}
question(#"כמה קשה הייתה היחידה עבורך?",
          "מה הייתה רמת הקושי של הלומדה עבורך?",
         type="single",
         answer("קלה"        ,correct=TRUE,  message = "תגובתך: קלה"),
         answer("בינונית"    ,correct=TRUE,  message = "תגובתך: בינונית"),
         answer("מאתגרת"     ,correct=TRUE,  message = "תגובתך: מאתגרת"),
         answer("קשה"        ,correct=TRUE,  message = "תגובתך: קשה"),
         answer("קשה מאוד"   ,correct=TRUE,  message = "תגובתך: קשה מאוד"),
         correct = "",
         incorrect = "",
         allow_retry = TRUE
)
```

```{r survey_q2, echo = FALSE}
question(#"עד כמה את/ה מרגיש/ה שהבנת את החומר שנלמד ביחידה?",
        "עד כמה החומר הנלמד ביחידה הועבר בצורה ברורה לדעתך?",
         type="single",
         answer("במידה רבה"      , correct=TRUE,  message = "תגובתך: במידה רבה"),
         answer("במידה סבירה" ,    correct=TRUE,  message = "תגובתך: במידה סבירה"),
         answer("באופן חלקי"  ,    correct=TRUE,  message = "תגובתך: באופן חלקי"),
         answer("במידה מעטה"  ,correct=TRUE,  message = "תגובתך: במידה מעטה"),
         answer("במידה מעטה מאוד",correct=TRUE,  message = "תגובתך: במידה מעטה מאוד"),
         correct = "",
         incorrect = "",
         allow_retry = TRUE
)
```

```{r survey_q3, echo = FALSE}
question("כמה זמן לקח לך לפתור את הלומדה בערך?",
         type="single",
         answer("פחות משעה"      , correct=TRUE,  message = "תגובתך: פחות משעה"),
         answer("שעה" ,    correct=TRUE,  message = "תגובתך: שעה"),
         answer("שעתיים"  ,    correct=TRUE,  message = "תגובתך: שעתיים "),
         answer("שלוש שעות"  ,correct=TRUE,  message = "תגובתך:  שלוש שעות"),
         answer("ארבע שעות ומעלה",correct=TRUE,  message = "תגובתך:  ארבע שעות ומעלה "),
         correct = "",
         incorrect = "",
         allow_retry = TRUE)
```

```{r survey_q4, echo = FALSE}
question("אילו נושאים בלומדה היו קשים במיוחד עבורך?",type="learnr_checkbox",
         answer("תיעוד נלווה", correct=TRUE),
         answer("בלוגים ופורומים אינטרנטיים", correct=TRUE),
         answer("שימוש נכון בבינה מלאכותית", correct=TRUE),
      
         answer("אף נושא", correct=TRUE),
         
         correct = "תגובתך התקבלה",
         try_again = "תגובתך התקבלה",
         incorrect = "תגובתך התקבלה",
         allow_retry = TRUE
)
```

```{r survey_q5, echo = FALSE}
  question_text(    "מצאת טעות? ספר/י לנו עליה (יש להקפיד ולציין את שם הפרק)",
   # "הערות ספציפיות לגבי חלקי היחידה השונים",
    incorrect = "תגובתך התקבלה",
    correct = "תגובתך התקבלה",
    answer(text = "/'ק/'ק",correct = TRUE),
    rows = 10,
    trim = FALSE
  )
```

```{r survey_q6, echo = FALSE}
  question_text(
        "יש לך עוד משהו לספר לנו? נשמח לשמוע משוב מפורט, הצעות כלליות או ספציפיות לפרק מסוים לגבי יחידה זאת.",
   # "הערות כלליות לגבי היחידה",
    incorrect = "תגובתך התקבלה",
    correct = "תגובתך התקבלה",
    answer(text = "/'ק/'ק",correct = TRUE),
    rows = 10,
    trim = FALSE
  )
```


## הגשה

עברו על הקובץ וודאו שהגשתם את כל התרגילים ועניתם על כל השאלות

במידה וכל התשובות שלכם תקינות יש ללחוץ על הכפתור: Generate, להעתיק את
הטקסט שמופיע בחלון למטה ולהגישו במודל\
בהצלחה!

```{r context="server"}
learnrhash::encoder_logic()
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui()
```


הדביקו את קוד הhash בתוך הגרשיים בפונקצייה הבאה כדי לוודא שעניתם על כל השאלות והתרגילים בלומדה.

```{r hash_check,  exercise = TRUE, exercise.eval = FALSE}

check_hash_code("")

```
